# CSCI 3155 Recitation 10
March 22, 2018

## Parametric Polymorphism

[Parametric polymorphism](https://en.wikipedia.org/wiki/Parametric_polymorphism) allows us to write reusable, type safe code. This is also called type parameters, generics, templates, or type variables.

As a reminder, we list type parameters in brackets similarly to how we list regular parameters in parentheses:
```scala
def f[T](x: T): Int = ???
```

## Exercise: Identity
Fix this identity function:
```scala
def id(x: Any): Any = x
```
So that the following operations compile:
```scala
val n: Int    = id(4)
val s: String = id("hi")
```

```scala
// BEGIN SOLUTION
def id[T](t: T): T = t
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36mid[39m
```

```scala
val n: Int    = id(4)
val s: String = id("hi")
```

Output:
```
[36mn[39m: [32mInt[39m = [32m4[39m
[36ms[39m: [32mString[39m = [32m"hi"[39m
```

## Exercise: Class
Improve the following class to take any type:
```scala
case class C(t: Int) {
    def get: Int = t
}
```

```scala
// BEGIN SOLUTION
case class C[T](t: T) {
    def get: T = t
}
// END SOLUTION
```

Output:
```
defined [32mclass[39m [36mC[39m
```

```scala
val x: Int = C(4).get
val y: String = C("4").get
```

Output:
```
[36mx[39m: [32mInt[39m = [32m4[39m
[36my[39m: [32mString[39m = [32m"4"[39m
```

### Exercise: Reusable data structures
Make a reusable tree data structure which can hold any arbitrary type at it's nodes, as long as all nodes in the tree have the same type. Use the constructors `EmptyPoly` and `NodePoly`

<!-- grade_id: cell-89ce853776b2aa84, points: 0, grade: false, solution: true -->
```scala
sealed trait TreePoly[T]
// BEGIN SOLUTION
case class EmptyPoly[T]() extends TreePoly[T]
case class NodePoly[T](l: TreePoly[T], d: T, r: TreePoly[T]) extends TreePoly[T]
// END SOLUTION
```

Output:
```
defined [32mtrait[39m [36mTreePoly[39m
defined [32mclass[39m [36mEmptyPoly[39m
defined [32mclass[39m [36mNodePoly[39m
```

```scala
// This should compile:
val t = NodePoly(
    NodePoly(
        EmptyPoly(),
        5,
        EmptyPoly()),
    6,
    EmptyPoly())
```

Output:
```
[36mt[39m: [32mNodePoly[39m[[32mInt[39m] = NodePoly(NodePoly(EmptyPoly(),5,EmptyPoly()),6,EmptyPoly())
```

```scala
// This should NOT compile:
NodePoly(
    NodePoly(
        EmptyPoly(),
        5,
        EmptyPoly()),
    "6",
    EmptyPoly())
```

Output:
```
cmd19.sc:2: type mismatch;
 found   : cmd19Wrapper.this.cmd17.wrapper.NodePoly[Int]
 required: cmd19Wrapper.this.cmd17.wrapper.TreePoly[Any]
Note: Int <: Any, but trait TreePoly is invariant in type T.
You may wish to define T as +T instead. (SLS 4.5)
    NodePoly(
            ^Compilation Failed
```

### Exercise: generic functions
Implement map, which should allow someone to take a `Tree[X]` and a function `X => Y` and get out a `Tree[Y]`.

<!-- grade_id: cell-11843a63b5b7990f, points: 0, grade: false, solution: true -->
```scala
// BEGIN SOLUTION
def map[X, Y](t: TreePoly[X], f: X => Y): TreePoly[Y] = t match {
    case EmptyPoly() =>
        EmptyPoly()
    case NodePoly(l, d, r) =>
        NodePoly(map(l, f), f(d), map(r, f))
}
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36mmap[39m
```

<!-- grade_id: cell-d725bf4fbe8683ac, points: 1, grade: true, solution: false -->
```scala
// TEST
val t2 = NodePoly(
    NodePoly(
        EmptyPoly(),
        "5",
        EmptyPoly()),
    "6",
    EmptyPoly())

assert(map(t, (i: Int) => i.toString) == t2)
// TEST
```

Output:
```
[36mt2[39m: [32mNodePoly[39m[[32mString[39m] = NodePoly(NodePoly(EmptyPoly(),5,EmptyPoly()),6,EmptyPoly())
```