# Tail Recursion

A recursive function is *tail-recursive* when the recursive subcall is the last computation of the recursive function body. The benefit of tail recursion is that it consumes less stack memory. See [here](https://docs.google.com/presentation/d/1q0l0RTgsTe2ux3VuhbsgeJZ50-rFeUhy-l8RjTiCiWc/edit?usp=sharing) for some visuals.

## Making a recursive function tail-recursive

A recursive function is *not* tail-recursive when it still has to perform computation after its recursive subcall returns. For example, in the following implementation of `factorial`, we need to multiply the returned result of `factorial(n-1)` by `n`.

```scala
def factorial(n : Int) : Int = {
    if (n == 0) {
        1
    } else {
        n * factorial(n - 1)
    }
}
```

To make this function tail-recursive, we want to take the work we perform *after* the recursive subcall and perform it *before* the recursive subcall. The trick is to perform that work on an accumulator variable, which then gets passed into the recursive subcall. When you reach the base case, you can simply return the work that you've accumulated in the accumulator variable.

The annotation `@tailrec` can be used to check if your function is properly tail recursive.

```scala
import scala.annotation.tailrec

// Notice the default argument for the accumulator (acc : Int = 1)

@tailrec
def factorial_tr(n : Int, acc : Int = 1) : Int = {
    if (n == 0) {
        acc
    } else {
        factorial_tr(n - 1, n * acc)
    }
}
```

The `tailrec` annotation causes a compile time error if the function cannot be optimized for tail recursion.

```scala
@tailrec
def factorial(n : Int) : Int = {
    if (n == 0) {
        1
    } else {
        n * factorial(n - 1)
    }
}
```

## Exercise: implement tail-recursive `reverse`

Here is a non-tail-recursive implementation of `reverse`.

```scala
def reverse(s : String) : String = {
    if (s.isEmpty()) {
        ""
    } else {
        val n = s.size
        s(n - 1) + reverse(s.slice(0, n - 1))
    }
}
```

```scala
assert(reverse("") == "", "Failed on \"\"")
assert(reverse("a") == "a", "Failed on \"a\"")
assert(reverse("abcde") == "edcba", "Failed on \"abcde\"")
assert(reverse("Mom's spaghetti") == "ittehgaps s'moM", "Failed on \"Mom's spaghetti\"")
```

Your task is to implement the tail-recursive version `reverse_tr`.

<!-- grade_id: cell-fcb94a97759a1e94, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
???
```

<!-- grade_id: cell-56099d8df8a74e27, points: 1, grade: true, solution: false -->
```scala
assert(reverse_tr("") == "", "Failed on \"\"")
assert(reverse_tr("a") == "a", "Failed on \"a\"")
assert(reverse_tr("abcde") == "edcba", "Failed on \"abcde\"")
assert(reverse_tr("Mom's spaghetti") == "ittehgaps s'moM", "Failed on \"Mom's spaghetti\"")
```

# Inductive definitions
## Grammars
## Exercise: Give grammar for inductively defined set of numbers

Convert the following inductive definition for sets given in english into a grammar.

A set ($S$) is defined inductively as follows:
- The empty set, $\emptyset$, is a set
- If $S$ is a set, then the insertion of a number is also a set: $S, n$.
- If $S$ is a set, then its complement is also a set: $\lnot S$
- If $S_1$ and $S_2$ are sets, then so are the following:
  - The intersection $S_1 \bigcap S_2$.
  - The union $S_1 \bigcup S_2$.
  - The subtraction $S_1 - S_2$.

See the following table for the constructor symbols which correspond to the given syntax (use these in your solution):

| Syntax              | Constructor    |
|---------------------|----------------|
| $n$                 | $\textbf{int}$ |
| $S$                 | $\textbf{Set}$ |
| $\emptyset$         | $Empty$        |
| $S, n$              | $Cons$         |
| $\lnot S$           | $Complement$   |
| $$S_1 \bigcap S_2$$ | $Intersection$ |
| $$S_1 \bigcup S_2$$ | $Union$        |
| $S_1 - S_2$         | $Subtraction$  |

$\LaTeX$ to get you started:
$$
\begin{array}{rcl}
    \textbf{Set} & \rightarrow & Cons(\textbf{Set}, n) \\
                 &           | & TODO \\
\end{array}
$$

## Exercise: Scala
Give the Scala case classes that represent the same structure.

<!-- grade_id: cell-08da72763665aaf9, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
???
```

Use the above case classes to create Scala objects representing the following:

1. $\emptyset$
2. $\{1, 2\}$
3. $\emptyset \bigcup \emptyset$
4. $(\emptyset \bigcup \emptyset) \bigcap (\lnot \emptyset)$
5. $\emptyset - \emptyset \bigcup \emptyset$

<!-- grade_id: cell-78c694043839ba98, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
???
```

```scala

```

The next section will not be covered during recitation, but has been left in the notebook for future reference.

# The Point of No Return
Until now, we have been calling `return` to pass values back from functions.  The `return` keyword is short hand for "skip the remaining code in this function".  Since `return` can cause unexpected behavior and make code less readable, we encourage students to avoid it from now on.

For example, consider the factorial function as it has been written in class.

```scala
def factorial_return(n : Int) : Int = {
    if (n != 0) {
        return n * factorial_return(n - 1)
    }
    return 1
}
factorial_return(5)
```

The "no return" version of this function is more than simply removing the `return` keyword.  We use an if *expression*.

```scala
def factorial(n : Int) : Int = {
    if (n != 0) {
        n * factorial(n - 1)
    } else {
        1
    }
}
factorial(5)
```

In many programming languages (e.g. Java, Python, and C++), we often use the `if` keyword to decide whether to execute a command.  If the \[condition\] is true, then the \[code\] runs. If the \[condition\] is false, nothing runs.

```
if([condition]){
    [code]
}
```

Scala evaluates `if` expressions to a value that may be stored in a variable, used as the return of a method, or more.
In the following code snippet, the first expression will be evaluated and stored in the result if the condition is true, otherwise the second expression will be evaluated and stored in result.
```
val result = if([condition]){ 
    [expression if true] 
} else { 
    [expression if false] 
}

```
Your task is to modify the `if` expression to fix the compilation error.

```scala
val n = 5
val b: String = if (n < 2) { "Less than or equal to two." } else { "Greater than or equal to two." }
```

Beware, an `if` without an `else` is rarely desirable.  If you simply remove the `return` keyword from the factorial function, it will always return `1`.

```scala
def factorial_bad(n : Int) : Int = {
    if (n != 0) {
        n * factorial_bad(n - 1)
    }
    1
}
factorial_bad(5)
```