```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment 7 - Covariance And Correlation, Independence - Sampling.ipynb")
```

# Homework 7: 

## Due on Gradescope


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Please see the **Course Syllabus for the Collaboration Policy**.

On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by an Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor.

We are here to help!  Visit OH Hours and/or post questions on Piazza!


If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced

```python

```

#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1abc | 3 | autograded
Question 1d | 2 | manual
Question 2 | 9 | manual
Question 3 | 7 | autograded
Question 4 | 2 | manual
Question 5ad | 5 | manual
Question 5bcef | 6 | autograded
Question 6abe| 7 | manual
Question 6cd| 3 | autograded
Question 7a | 1 | autograded
Question 7bcd | 5 | manual
|Total | 50

```python
import numpy as np
import pandas as pd
import math
import scipy.stats as stats
from scipy import special
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
from IPython.display import display, Latex, Markdown

plt.style.use('fivethirtyeight')
# Default plot configurations
plt.rcParams['figure.figsize'] = (16,8)
plt.rcParams['figure.dpi'] = 150
sns.set()

%matplotlib inline
```

<a id='top'></a>
---
**Shortcuts:**  [Problem 2](#p2) | [Problem 3](#p3) | [Problem 4](#p4) |  [Problem 5](#p5) | [Problem 6](#p6) | [Problem 7](#p7) | 

               
---
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />



# Part 1:  Joint Distributions, Covariance and Correlation

## (5 pts) Problem 1 :

Suppose $X$ and $Y$ are two discrete random variables with the following joint probability mass function:



$$
\begin{array}{|c|c|c|}
\hline
 & Y=1 & Y=2 \\
\hline
X=1 & \frac{1}{3} &  \\
\hline
X=2 & \frac{1}{6} & 0  \\
\hline
X=4 & \frac{1}{12} & \frac{1}{3}  \\
\hline
\end{array}
$$




**Part A)**
What is $P(X=1, Y=2)$?  


```python
prob_xequals1_and_yequals2 = ...
```

```python
grader.check("q1a")
```

**Part B)**
Find the marginal PMF for X, (i.e. find $P(X=x)$ for $x=1, 2, 4$).  

```python
prob_x_equals_1 = ...
prob_x_equals_2 = ...
prob_x_equals_4 = ...
```

```python
grader.check("q1b")
```

**Part C)**
Find the marginal PMF for Y (i.e. find $P(Y=y)$ for $y = 1, 2$)

```python
prob_y_equals_1 = ...
prob_y_equals_2 = ...
```

```python
grader.check("q1c")
```

<!-- BEGIN QUESTION -->

**Part D)**
Are $X$ and $Y$ independent or dependent?  Fully justify your answer in the cell below using LaTeX and the mathematical definition of independence.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer to question 1d, enter a boolean value, True or False (note to receive credit on 1d you must mathematically justify your answer in the cell above). 

```python
x_y_independent = ...
```

```python
grader.check("q1d_check")
```

***

[Back to top](#top)

<a id='p2'></a>


## (9 pts) Problem 2 :

The following parts are **NOT** related:


<!-- BEGIN QUESTION -->

**Part A)**
If Cov$(X,Y)=0$, what does this tell us about the random variables X and Y?  


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Part B)**  Given the following joint pmf for discrete random variables $X$ and $Y$:

$$
\begin{array}{|c|c|c|c|}
\hline
 & Y=0 & Y=1 & Y=2 \\
\hline
X=0 & \frac{1}{6} & \frac{1}{4} &\frac{1}{8} \\
\hline
X=1 & \frac{1}{8} & \frac{1}{6} &\frac{1}{6} \\
\hline
\end{array}
$$


 - i).  Calculate $Cov(X,Y)$.  
 - ii).  Calculate $\rho(X,Y)$   
 


Show all steps for both parts using Markdown and LaTeX in the cell below:

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for Problem 2bi, enter the answer you came up with (just the number) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the cell above.  

```python
cov_xy = ...
```

```python
grader.check("q2bi_check")
```

To check your answer for Problem 2bii, enter the answer you came up with (just the number) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the cell above.  

```python
p_xy = ...
```

```python
grader.check("q2bii_check")
```

<!-- BEGIN QUESTION -->

**Part C)**
This part is **NOT** related to the parts above.  
Suppose you're only given the following information about two joint random variables $X$ and $Y$:  

$$\mu_X = 6, \hspace{2mm} \mu_Y = 5, \hspace{2mm} \sigma_X^2 = 4, \hspace{2mm} \sigma_Y^2=9 \text{ and }  E[XY]=27$$.  

For each of the quantities below, calculate if you have enough information, showing all steps.  If not, explain what additional info you'd need.

i).  $Cov(X,Y)$

ii).  $Cov(Y,X)$

iii).  $\rho(X,Y)$

Answer all parts in the ONE markdown cell below, fully justifying your answer:

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answers for Problem 2c, enter the answers you came up with (just the numbers) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the cell above.  

```python
cov_xy = ...
cov_yx = ...
p_xy = ...
```

```python
grader.check("q2c_check")
```

<hr>

# Part 2 Introduction

The outcome of the US presidential election in 2016 took many people and many pollsters by surprise. In this assignment we will carry out a simulation study / post mortem in an attempt to understand what happened.

***

[Back to top](#top)

<a id='p3'></a>


## (7 pts) Problem 3

In Pennsylvania, 6,165,478 people voted in the 2016 Presidential election.
Trump received 48.18% of the vote and Clinton recieved 47.46%.
This doesn't add up to 100% because other candidates received votes.
All together these other candidates received 100% - 48.18% - 47.46% = 4.36% of the vote.

The table below displays the counts and proportions.


|   Voted for   |  Trump|    Clinton|    Other|
|-----------|-----------|-----------|---------|
| Probability      |   0.4818   | 0.4746  |   0.0436 |
| Number of people | 2,970,733  | 2,926,441 | 268,304 |

### 3a)
Suppose we pick a simple random sample of 20 of the 6,165,478 Pennsylvania (PA) voters. 

In the sample, let $N_T$ be the number of Trump voters, $N_C$ the number of Clinton voters, and $N_O$ the number of "other" voters. Then $N_T$, $N_C$, and $N_O$ are random: they depend on how the sample comes out. Such quantities are called "statistics".

Pick the correct option: $N_T + N_C + N_O$ is equal to

(a) 3

(b) 20

(c) 6,165,478

(d) a random quantity

*Note*: if your answer is (d), put  **ans_3a = 'd'**  for the purpose of grading. Passing the test below **does NOT** mean you answer is correct (there are hidden tests that will check this question when you submit your assignment). The test below only checks if your answer is in the correct format.

```python
ans_3a = ...
```

```python
grader.check("q3a")
```

### 3b) 
Pick the correct option.

A simple random sample of 20 PA voters is like a sample drawn at random with replacement, because

(a) that's the definition of "simple random sample"

(b) there are only 3 categories of voters, which is small in comparison to 20

(c) there are only 20 people in the sample, which is small in comparison to the total number of PA voters

(d) all PA voters are equally likely to be selected

*Note*: if your answer is (d), put **ans_3a = 'd'**  for the purpose of grading. Passing the test below **does NOT** mean you answer is correct (there are hidden tests that will check this question when you submit your assignment). The test below only checks if your answer is in the correct format.

```python
ans_3b = ...
```

```python
grader.check("q3b")
```

### 3c)

Let $t$, $c$, and $o$ be any three non-negative integers. For the sample defined in Part **a**, the probability that the sample contains $t$ Trump voters, $c$ Clinton voters and $o$ "other" voters is denoted by $P(N_T = t, N_C = c, N_O = o)$.

Define a function `prob_sample_counts` that takes any three non-negative integers $t$, $c$, and $o$, and returns $P(N_T = t, N_C = c, N_O = o)$. The probability is 0 for some choices of the arguments, and your function should return 0 in those cases.

Remember that Part **b** implies you can use results for sampling with replacement.


```python
def prob_sample_counts(t, c, o):
    """
    Input:
    t - number of votes for Trump
    c - number of votes for Clinton
    o - number of votes for Other
    Return: The probability of getting a such sample
    """
    ...
    # your code above this line
```

```python
grader.check("q3c")
```

### 3d) 

Check that your function determines a probability distribution by summing all the positive terms. You know what the sum should be. **Hint: You might need to iterate through certain values of c and t to compute the sum of all positive probabilities.**

```python
prob_3d = ...
prob_3d
```

```python
grader.check("q3d")
```

### 3e)
Use `prob_sample_counts` to find the chance that the sample consists of 11 Trump voters, 8 Clinton voters, and 1 "other" voter.

```python
prob_3e = ...
prob_3e
```

```python
grader.check("q3e")
```

### 3f)
Write an expression that evaluates the chance that the sample contains at least 7 Trump voters, at least 7 Clinton voters, and at least 1 "other" voter. **Hint: You might find 3d useful.**

```python
prob_3f = ...
prob_3f
```

```python
a=np.array([prob_sample_counts(t, c, 20-(t+c)) for t in range(7, 13) for c in range(7, 20-t)])

print(a)
```

```python
grader.check("q3f")
```

## Election Polling

Political polling is a type of public opinion polling that can at best represent a snapshot of public opinion at the particular moment in time. Voter opinion shifts from week to week, even day to day, as candidates battle it out on the campaign field.

Polls usually start with a "horse-race" question, where respondents are asked whom they would vote for in a head-to-head race if the election were tomorrow: Candidate A or Candidate B. The survey begins with this question so that the respondent is not influenced by any of the other questions asked in the survey. Some of these other questions are asked to help assess how likely is it that the respondent will vote. Other questions are asked about age, education, and sex in order to adjust the findings if one group appears overly represented in the sample.

Pollsters typically use [random digit dialing](https://en.wikipedia.org/wiki/Random_digit_dialing) to contact people.

<!-- BEGIN QUESTION -->

***

[Back to top](#top)

<a id='p4'></a>

## (2 pts) Problem 4


If we're trying to predict the results of the Clinton vs. Trump 2016 presidential race:

i).  What is the population of interest? 

ii).  What is the sampling frame?

Give both of your answers in the same below in Markdown.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### How might the sampling frame differ from the population?

After the fact, many experts have studied the 2016 election results. For example, according to the American Association for Public Opinion Research (AAPOR), predictions made before the election were flawed for three key reasons:

1. voters changed their preferences a few days before the election
2. those sampled were not representative of the voting population, e.g., some said that there was an overrepresentation of college graduates in some poll samples 
3. voters kept their support for Trump to themselves (hidden from the pollsters)



The 1st and 3rd biases (voters changing preference and voters hiding their preference) can only be determined after the election. It might be possible to assess this with a post election survey where people were asked about these things.   

In the rest of this assignment we will focus on analyzing the 2nd  type of bias:


+ Question 5: We will carry out a study of the sampling error when there is no bias. In other words, we will try to compute the chance that we get the election result wrong even if we collect our sample in a manner that is completely correct. In this case, any **failure of our prediction is due entirely to random chance**.
+ Question 6: We will carry out a study of the sampling error when there is bias of the second type from the list above. In other words, we will try to compute the chance that we get the election result wrong if we have a small systematic bias. In this case, any failure of our prediction is due to a **combination of random chance and our bias**.



### How large was the sampling error?

In some states the race was very close, and it may have been simply sampling error, i.e., random chance that the majority of the voters chosen for the sample voted for Clinton.

A 2- or 3-point polling error in Trump’s favor (typical error historically) would likely be enough to tip the Electoral College to him.

One year after the 2016 election, Nate Silver (https://substack.com/@natesilver) wrote in
*The Media Has A Probability Problem* that
the "media’s demand for certainty -- and its lack of statistical rigor -- is a bad match for our complex world."
FiveThirtyEight forecasted that Clinton had about a 70 percent chance of winning.  

We will first carry out a simulation study to assess the impact of the sampling error on the predictions.


## The Electoral College

The US president is chosen by the Electoral College, not by the
popular vote. Each state is alotted a certain number of 
electoral college votes, as a function of their population.
Whomever wins in the state gets all of the electoral college votes for that state.

There are 538 electoral college votes (hence the name of the website Nate Silver started, FiveThirtyEight: https://projects.fivethirtyeight.com/polls/).

Pollsters correctly predicted the election outcome in 46 of the 50 states. 
For these 46 states Trump received 231 and Clinton received 232 electoral college votes.

The remaining 4 states accounted for a total of 75 votes, and 
whichever candidate received the majority of the electoral college votes in these states would win the election. 

These states were Florida, Michigan, Pennsylvania, and Wisconsin.

|State |Electoral College Votes|
| --- | --- |
|florida | 29 |
|michigan | 16 |
|pennsylvania | 20 |
|wisconsin | 10|

For Donald Trump to win the election, he had to win either:
* Florida + one (or more) other states
* Michigan, Pennsylvania, and Wisconsin


In the final results the electoral margins were very narrow in these four states, as seen below:


|State | % Trump |   % Clinton | Total Voters |
| --- | --- |  --- |  --- |
|florida | 49.02 | 47.82 | 9,419,886  | 
|michigan | 47.50 | 47.27  |  4,799,284|
|pennsylvania | 48.18 | 47.46 |  6,165,478|
|wisconsin | 47.22 | 46.45  |  2,976,150|

Those narrow electoral margins can make it hard to predict the outcome **given the sample sizes** that the polls used. 

---
## Simulation Study of the Sampling Error

Now that we know how people actually voted, we can carry
out a simulation study that imitates the polling.

Our ultimate goal in this problem is to use the actual voting percentages to simulate polling results and **understand the chance that we will incorrectly call the election for Hillary Clinton even if our sample was collected with absolutely no bias**.

<!-- BEGIN QUESTION -->

***

[Back to top](#top)

<a id='p5'></a>


## Problem 5 (11 pts) 

#### Part A

For your convenience, the actual results of the vote in the four pivotal states is repeated below:

|State |% Trump |   % Clinton | Total Voters |
| --- | --- |  --- |  --- |
|florida | 49.02 | 47.82 | 9,419,886  | 
|michigan | 47.50 | 47.27  |  4,799,284|
|pennsylvania | 48.18 | 47.46 |  6,165,478|
|wisconsin | 47.22 | 46.45  |  2,976,150|


Using the table above, write a function `draw_state_sample(N, state)` that returns a sample with replacement of N voters from the given state, using the percentages given in the table above. Your result should be returned as a list, where the first element is the number of Trump votes, the second element is the number of Clinton votes, and the third is the number of Other votes. For example, `draw_state_sample(1500, "florida")` could return `[727, 692, 81]`. You may assume that the state name is given in all lower case.  

**Hint:** You might find `np.random.multinomial` useful.


```python
def draw_state_sample(N, state):
    ...
    
    

```

```python
grader.check("q5a")
```

<!-- END QUESTION -->

#### Part B

Now, create a function `trump_advantage` that takes in list of sample of votes (like the one returned by `draw_state_sample`) and returns the difference in the proportion of votes between Trump and Clinton. For example `trump_advantage([100, 60, 40])` should return `0.2`, since Trump had 50% of the votes in this sample and Clinton had 30%.  In contrast, `trump_advantage([60, 100, 40])` should return `-0.2`

```python
def trump_advantage(voter_sample):
    ...
```

```python
grader.check("q5b")
```

#### Part C

Simulate Trump's advantage across 100,000 simple random samples of 1500 voters for the state of Pennsylvania and store the results of each simulation in an `np.array` called `simulations`. 

That is, `simulations[i]` should be Trump's proportion advantage for the `i+1`th simple random sample.

```python
simulations = ...

```

```python
grader.check("q5c")
```

<!-- BEGIN QUESTION -->

#### Part D

i).  Make a **frequency** histogram of `simulations`.   This is a histogram of the sampling distribution of Trump's proportion advantage in Pennsylvania. 

Hint: You should use the [`plt.hist`](https://matplotlib.org/3.1.1/api/_as_gen/matplotlib.pyplot.hist.html) function in your code.

Make sure to include a title as well as axis labels. You can do this using `plt.title`, `plt.xlabel`, and `plt.ylabel`.

ii).  Based on your simulation, what is the probability that a random sample of 1500 will correctly predict that Trump wins Pennsylvania? (i.e. what proportion of these simulations predict a Trump victory?)  Assign your answer to `prob_penn_1500_random_correct`


```python
# Part (i):
...
# your code for the histogram above here.  The code below plots a red marker at the mean:
plt.scatter(simulations.mean(), -1, marker='^', color='red', s=500)
```

```python
# Part (ii):
prob_penn_1500_random_correct = ...

prob_penn_1500_random_correct
```

<!-- END QUESTION -->

Your result above allows us to calculate the chance that we will incorrectly call the election for Hillary Clinton in Pennsylvania based on the results of this random sample of 1500 even if our sample was collected with absolutely no bias:


```python
prob_penn_1500_random_incorrect = 1-prob_penn_1500_random_correct
prob_penn_1500_random_incorrect
```

#### Part E

Now write a function `trump_wins(N)` that creates a random sample of N voters for **each of the four crucial states** (that is N voters from Florida, N voters from Michigan, N voters from Pennsylvania, and N voters from Wisconsin) and returns 1 if Trump is predicted to win based on these samples and 0 if Trump is predicted to lose.   
(Hint: use `draw_state_sample` from above) 

Recall that for Trump to win the election, he must either:
* Win the state of Florida and 1 or more other states
* Win Michigan, Pennsylvania, and Wisconsin

Then repeat 100,000 simulations of the election, i.e. call `trump_wins(1500)` 100,000 times and calculate what proportion of these simulations predict a Trump victory. Give your answer as a float stored in `proportion_trump`.

This number represents the percent chance that samples of 1500 from each of the 4 states will correctly predict Trump's victory *if the samples were collected with absolutely no bias*. 


```python
def trump_wins(N):
    ...
    
proportion_trump = ...
proportion_trump
```

```python
grader.check("q5e")
```

#### Part F
Based on your work above, what is the chance that we will incorrectly call the election for Hillary Clinton based on the results of random samples of 1500 people in these 4 battleground states **even if our samples were collected with absolutely no bias?**    

```python
prob_election_1500_random_incorrect = ...
prob_election_1500_random_incorrect
```

```python
grader.check("q5f")
```

---
## Simulation Study of Selection Bias

According to [an article](https://blogs.lse.ac.uk/usappblog/2018/02/01/better-poll-sampling-would-have-cast-more-doubt-on-the-potential-for-hillary-clinton-to-win-the-2016-election/#Author) by Grotenhuis, Subramanian, Nieuwenhuis, Pelzer and Eisinga:

> In a perfect world, polls sample from the population of voters, who would state their political preference perfectly clearly and then vote accordingly.

That's the simulation study that we just performed. 


It's difficult to control for every source of selection bias.
And, it's not possible to control for some of the other sources of bias.

Next we investigate the effect of small sampling bias on the polling results in these four battleground states.  

Throughout this problem, we'll examine the impacts of a 0.5 percent bias in favor of Clinton in each state. Such a bias has been suggested because highly educated voters tend to be more willing to participate in polls.

<!-- BEGIN QUESTION -->

***

[Back to top](#top)

<a id='p6'></a>



## Problem 6 (10 pts)

Throughout this problem, adjust the selection of voters so that there is a 0.5% bias in favor of Clinton in each of these states. 

For example, in Pennsylvania, Clinton received 47.46\% of the votes and Trump 48.18\%. Increase the population of Clinton voters to 47.46\% + 0.5\%  and correspondingly decrease the percent of Trump voters. 

#### Part A

Simulate Trump's advantage across 100,000 simple random samples of 1500 voters for the **state of Pennsylvania** and store the results of each simulation in an `np.array` called `biased_simulations`.

That is, `biased_simulation[i]` should hold the result of the `i+1`th simulation.

That is, your answer to this problem should be just like your answer from Question 5C, but now using samples that are biased as described above.


```python
def draw_biased_state_sample(N, state):
    
    ...
    
biased_simulations = ...
```

```python
grader.check("q6a")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part B

Create a plot of **overlaid DENSITY** histograms of the following:
 - The new sampling distribution of Trump's proportion advantage in Pennsylvania using these biased samples 
  - The sampling distribution of the unbiased samples from Problem 5D (plotted as a density, not a frequency histogram)
  
Include 2 markers (of different colors) with the sample means for each distribution (see 5D for code how to do this).  The colors of the markers should correspond to the colors of the density histograms.

Make sure to give your plot a title, label the x and y axes and include a legend.
Use the parameter `alpha` to adjust the transparency of each histogram.




```python
...
```

<!-- END QUESTION -->

#### Part C

Use your result above to calculate the chance that we will incorrectly call the election for Hillary Clinton in Pennsylvania based on the results of this random sample of 1500 if our sample was collected with this $0.5\%$ bias:


```python
prob_penn_1500_biased_incorrect = ...
prob_penn_1500_biased_incorrect
```

```python
grader.check("q6c")
```

#### Part D

Now perform 100,000 simulations of $1500$ voters from **each of the four states** (1500 voters from Pennsylvania, 1500 voters from Michigan, etc)  and return the proportion of these simulations that result in a Trump victory. This is the same fraction that you computed in Question 5e, but now using your biased samples.

Give your answer as `proportion_trump_biased`.

This number represents the chance that a sample biased 0.5% in Hillary Clinton's favor will correctly predict Trump's victory. 

```python
def trump_wins_biased(N): 
    ...

proportion_trump_biased = ...
proportion_trump_biased
```

```python
grader.check("q6d")
```

#### Part E

<!-- BEGIN QUESTION -->

Summarize the findings from these simulations:

i).  Based on your simulations, what was the **chance of error** in correctly predicting that Trump wins using the **unbiased** samples of $1500$ people from each state?      Many people, even well educated ones, assume that this number should be $0\%$. After all, how could a non-biased sample be wrong?   Give a mathematical explanation as to why it isn't $0\%$ (or close to $0\%$).   This is the type of incredibly important intuition we hope to develop in you throughout this class and your future data science coursework.


ii).  What was the chance of error in predicting the results using the **biased** samples and how different is it from your answer in part(i)?  Recall, we only biased the samples by $0.5\%$. However, even a bias this small in the percentages can lead to a much larger chance of error in prediction of the final result.

_Type your answer here, replacing this text._

<!-- END QUESTION -->


***

[Back to top](#top)

<a id='p7'></a>

## Question 7 (6 pts)

Would increasing the sample size have helped?

#### Part A


Try a sample size of 5,000 (in each of the 4 states) and run 100,000 simulations of a sample with replacement. What proportion of the 100,000 times is Trump incorrectly predicted to lose the election in the unbiased setting? In the biased setting?

Give your answers as `high_sample_size_unbiased_proportion_trump` and `high_sample_size_biased_proportion_trump`.



*Note: Don't worry if your solution takes a while to run.* 

```python
high_sample_size_unbiased_proportion_incorrect = ...
high_sample_size_biased_proportion_incorrect = ...

print(high_sample_size_unbiased_proportion_incorrect, high_sample_size_biased_proportion_incorrect)

```

```python
grader.check("q7a")
```

<!-- BEGIN QUESTION -->

#### Part B

Compare your observations from 7a to your observations in 6d.  Did the chance of error increase or decrease in each case and why?  What do these changes imply about the impact of sample size on the sampling error and on the bias?   

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part C

Is it possible to correctly predict Trump's victory with less than $1\%$ error using **unbiased sampling?**  Rerun the simulation (in each of the 4 states) with increasing sample sizes and 100,000 simulations to determine if you can find an approximate minimum sample size (it doesn't have to be exact) such that the probability of correctly predicting Trump's victory is at least $99\%$ (assuming your sample is unbiased).  



```python
...
# your code above this line.
# output the number of samples you used to get to at least 99% accuracy.
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part D

Is it possible to correctly predict Trump's victory with less than $1\%$ error using **biased sampling?**  Use the code cell below to rerun the simulation (in each of the 4 states) with increasing sample sizes.  What happens to the probability of error?  Explain in the markdown cell below.



_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
print("10000 runs gives", np.mean([trump_wins_biased(10000) for i in range(100000)]))
print("20000 runs gives",np.mean([trump_wins_biased(20000) for i in range(100000)]))
print("30000 runs gives",np.mean([trump_wins_biased(30000) for i in range(100000)]))
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 7!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 7 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW7 Autograded assignment in Gradescope.

* **Homework 7 Manually Graded**: Submit your hw07.PDF to the HW7 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw07.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 1)
```
>>> # TEST
>>> assert np.isclose(prob_xequals1_and_yequals2, 1/12)
```


Tests for q1b (points: 1)
```
>>> # TEST
>>> assert np.isclose(prob_x_equals_1, 5/12)
```
```
>>> # TEST
>>> assert np.isclose(prob_x_equals_2, 1/6)
```
```
>>> # TEST
>>> assert np.isclose(prob_x_equals_4, 5/12)
```


Tests for q1c (points: 1)
```
>>> # TEST
>>> assert np.isclose(prob_y_equals_1, 7/12)
```
```
>>> # TEST
>>> assert np.isclose(prob_y_equals_2, 5/12)
```


Tests for q1d_check (points: 0)
```
>>> # TEST
>>> assert x_y_independent == False
```


Tests for q2bi_check (points: 0)
```
>>> # TEST
>>> assert np.isclose(cov_xy, 1/24)
```


Tests for q2bii_check (points: 0)
```
>>> # TEST
>>> assert np.isclose(p_xy, 2*np.sqrt(3/1001))
```


Tests for q2c_check (points: 0)
```
>>> # TEST
>>> assert cov_xy == -3
```
```
>>> # TEST
>>> assert cov_yx == -3
```
```
>>> # TEST
>>> assert p_xy == -1/2
```


Tests for q3a (points: [0])
```
>>> # TEST
>>> assert ans_3a in ['a', 'b', 'c', 'd']
```


Tests for q3b (points: [0])
```
>>> # TEST
>>> assert ans_3b in ['a', 'b', 'c', 'd']
```


Tests for q3c (points: 2)
```
>>> # TEST
>>> assert (0 <= prob_sample_counts(2, 5, 13) <= 1)
```
```
>>> # TEST
>>> assert (0 <= prob_sample_counts(5, 1, 14) <= 1)
```
```
>>> # TEST
>>> assert (0 <= prob_sample_counts(10, 8, 2) <= 1)
```
```
>>> # TEST
>>> assert (0 <= prob_sample_counts(6, 12, 2) <= 1)
```
```
>>> # TEST
>>> assert np.isclose(prob_sample_counts(6, 12, 3), 0, atol = 1e-5)
```
```
>>> # TEST
>>> assert np.isclose(prob_sample_counts(20, 1, 3), 0, atol = 1e-12)
```
```
>>> # TEST
>>> assert np.isclose(prob_sample_counts(6, 6, 2), 0, atol = 1e-5)
```
```
>>> # TEST
>>> assert np.isclose(prob_sample_counts(6, 12, 2), 0.01095, atol = 1e-4)
```
```
>>> # TEST
>>> assert np.isclose(prob_sample_counts(13, 7, 0), 0.03169, atol = 1e-4)
```
```
>>> # TEST
>>> assert np.isclose(prob_sample_counts(13, 5, 2), 0.00562, atol = 1e-4)
```


Tests for q3d (points: 1)
```
>>> # TEST
>>> assert np.isclose(prob_3d, 1, atol = 1e-4)
```


Tests for q3e (points: 1)
```
>>> # TEST
>>> assert np.isclose(prob_3e, 0.05509, atol = 1e-4)
```


Tests for q3f (points: 1)
```
>>> # TEST
>>> assert np.isclose(prob_3f, 0.46939, atol = 1e-4)
```


Tests for q5a (points: 0)
```
>>> # TEST
>>> assert len(draw_state_sample(1500, "florida")) == 3
```
```
>>> # TEST 
>>> assert sum(draw_state_sample(1500, "michigan")) == 1500
```
```
>>> # TEST
>>> q5a_penn = draw_state_sample(1500, "pennsylvania")
>>> trump_win_penn = (q5a_penn[0] - q5a_penn[1]) / 1500
>>> assert abs(trump_win_penn - 0.007) <= 0.12
```


Tests for q5b (points: [0, 0.5, 0.5])
```
>>> # TEST
>>> assert (-1 < trump_advantage(draw_state_sample(1500, "wisconsin")) < 1)
```
```
>>> # TEST
>>> assert np.isclose(trump_advantage([100, 60, 40]), 0.2)
```
```
>>> # TEST
>>> assert np.isclose(trump_advantage([10, 30, 10]), -0.4)
```


Tests for q5c (points: [0, 0, 2])
```
>>> # TEST
>>> assert len(simulations) == 100000
```
```
>>> # TEST
>>> assert sum([-1 < x < 1 for x in simulations]) == len(simulations)
```
```
>>> # TEST
>>> assert abs(np.mean(simulations) - 0.007) <= 0.016
```


Tests for q5e (points: [0, 2])
```
>>> # TEST
>>> trump_wins(1000) in [0, 1]
True
```
```
>>> # TEST
>>> assert abs(proportion_trump - 0.695) <= 0.02
```


Tests for q5f (points: 1)
```
>>> # TEST
>>> 
>>> assert abs((1-proportion_trump) - .3)<= 0.03
```


Tests for q6a (points: 0)
```
>>> # TEST
>>> assert len(draw_biased_state_sample(1000, "wisconsin")) == 3
```
```
>>> # TEST
>>> assert sum(draw_biased_state_sample(1000, "michigan")) == 1000
```
```
>>> # TEST
>>> q6_1 = draw_biased_state_sample(3000, "florida")
>>> assert (q6_1[0] > q6_1[2] and q6_1[1] > q6_1[2])
```
```
>>> # TEST
>>> assert len(biased_simulations) == 100000
```
```
>>> # TEST
>>> assert sum([-1 < x < 1 for x in biased_simulations]) == len(biased_simulations)
```
```
>>> # TEST
>>> assert abs(np.mean(biased_simulations) + 0.003) <= 0.12
```


Tests for q6c (points: 1)
```
>>> assert abs(prob_penn_1500_biased_incorrect - 0.55) <= 0.025
```


Tests for q6d (points: 2)
```
>>> # TEST
>>> assert abs(proportion_trump_biased - 0.466) <= 0.04
```


Tests for q7a (points: 1)
```
>>> # TEST
>>> np.abs(high_sample_size_unbiased_proportion_incorrect - 0.171) <= 0.02
True
```
```
>>> # TEST 
>>> np.abs(high_sample_size_biased_proportion_incorrect - 0.555) <= 0.04
True
```
