```python
# Initialize Otter
import otter
grader = otter.Notebook("HW2.ipynb")
```

### Grading
The final score that you will receive for your programming assignment is generated in relation to the total points set in your programming assignment item—not the total point value in the autograder output.     
Autograded quesions may have hidden tests and/or public tests. You can see the public test results when you run the notebook with otter grader. Usually public tests are for simple checks (such as variable types) and may not indicate that the answer is correct.     
**DO NOT CHANGE VARIABLE OR METHOD SIGNATURES** The autograder will not work properly if your change the variable or method signatures. 

# Logistic Regression

```python
# importing all the required libraries

from math import exp
import numpy as np
import pandas as pd
from sklearn.linear_model import LogisticRegression
from sklearn.datasets import load_breast_cancer
from sklearn.model_selection import train_test_split
import matplotlib.pyplot as plt
%matplotlib inline

from sklearn.metrics import roc_auc_score
from sklearn.metrics import roc_curve

```

## Part A. Understanding Logistic Function

### Q1. [5 points] 
Your first task is to complete the function `gen_logistic` in the following cell so as to be able to generate the logistic function for a given input. The logistic function is a type of <em>sigmoid</em> function which has an 'S'-shape and 'squashes' its inputs to a value lying in the range [0,1]. Other sigmoid functions include the hyperbolic-tangent funcition (`tanh(x)`) and the error function (`erf(x)`).[https://en.wikipedia.org/wiki/Sigmoid_function].     
Hint: In the logistic function, the variable of the logistic function is the logit $z = w \cdot x +b$.

<!-- grade_id: cell-29a60883217c420b, points: 0, grade: false, solution: true -->
```python
def gen_logistic(x, w=1, b=0):
    """
    outputing the logistic output for an input x
    :param x: scalar or numpy array of shape (n_samples, n_features). If only one feature, it must have the shape of (n_samples,1).
    :param w: weight(s); either scalar or numpy array of shape (1, n_features)
    :param b: bias; either scalar or numpy array of shape (1,)
    returns y of shape (n_samples,)
    """    
    # TODO: Finish this function to return the output of applying the sigmoid
    # function to the input x (Please do not use external libraries) store 
    # the output in y and return y. Do not change the default parameter values.
    # Hint: This function will be used in any input shape scalar (0d), 1d vector, and 2d arrays. Please make sure it can handle all those. Following reshaping codes might help.
    # Hint2: You may use design matrix using concatenation (e.g. np.concat), but it is not necesary.
    
    y =0 
    if np.isscalar(x):
        x = np.array(x).reshape((1,1))
    if np.isscalar(w):
        w = np.array(w).reshape((1,1))
    if np.isscalar(b):
        b = np.array(b).reshape((1,1))  
    if b.shape==(1,):
        b= b.reshape((1,1))  

    ...
    
    return y.reshape(y.shape[0],)
```

<!-- BEGIN QUESTION -->

### Q2. Role of parameters in a logistic function [10 pts] 
**2-a) [5 points]:** Generate a vector x of length N with values lying between limits Xa and Xb (for this you will have to choose your own limits; play around with different values) and apply the `gen_logistic` function to this vector.  Proceed to plot the output and verify the shape of the output. If your decision boundary value is about the center of your x range, you will see an S-shape. Your final plot should show the S-curve. 

<!-- grade_id: cell-6eb2c8a94fdaff49, points: 0, grade: false, solution: true -->
```python
# TODO: change the values of N, a and b below to check how the output of your function works
# Use a value for N greater than 1 and any limits a and b so that an S-shape graph is generated

N = 1
Xa = 0
Xb = 1
w = 1
b = 0



x = np.expand_dims(np.linspace(Xa,Xb,N), axis=1)
y = gen_logistic(x, w, b)

fig, ax = plt.subplots(nrows=1,ncols=1,figsize=(12,7))
ax.plot(x,y, lw=2)
ax.set_xlabel("x", fontsize=16)
ax.set_ylabel("y", fontsize=16)
ax.set_title("Logistic/Sigmoid Function", fontsize=16)

...

```

<!-- END QUESTION -->

###  Answer following questions
#### 2-b) Increasing w will make the curve transition sharply: [2 pts, True/False]

<!-- grade_id: cell-cba612162fd69391, points: 0, grade: false, solution: true -->
```python
# uncomment sharp_transition and answer qustion 1. above 
# replace string with 'True' or 'False' 
# sharp_transition = ''
...
```

#### 2-c) If b increases by 1, then the decision boundary x decreases by 1: [1 pts, True/False]

<!-- grade_id: cell-b6f28ec800691431, points: 0, grade: false, solution: true -->
```python
# uncomment x_decreases_by_1 and answer question 2. above
# replace string with 'True' or 'False' 
# x_decreases_by_1 = ''
...
```

#### 2-d) If $w$'s sign flips (and no change in the non-zero-valued $b$), the logistic curve flips vertically without changing the boundary threshold's position. [2 pts, True/False]

```python
# uncomment vertical_flip and answer question 3. above
# replace string with 'True' or 'False' 
# vertical_flip = ''

...
```

# Part B. Logistic regression using Scikit-learn

### Q3.  [10 pts]
Performing binary classification using logistic regression on the breast-cancer dataset. In this part you will be exposed to different methods within the scikit-learn LogisticRegression class so you can build a classifier. In this part, we will use breast cancer data from sklean dataset library.

#### 3-a)  Complete below codes to split train and test data  [5 pts]

<!-- grade_id: cell-10976d76d3a81f36, points: 0, grade: false, solution: true -->
```python
# Importing the breast-cancer dataset from sklearn datasets

class BC_data:
    """
    class to import the breast cancer dataset from sklearn
    
    """
    def __init__(self):
        
        x, y = load_breast_cancer(return_X_y= True)
        self.x_train = None 
        self.x_test = None 
        self.y_train = None 
        self.y_test = None
        
        #TODO: Split the data into training and test data (use train_test_split sklearn) 
        # such that the test data size is 25% of total number of observations
        # No need to rescale the data. Use the data as is.
        
        ...
        
        
data = BC_data()
```

#### 3-b) Build and Fit Logistic Regression Model [5 pts]

<!-- grade_id: cell-7c7287ce51bbec9b, points: 0, grade: false, solution: true -->
```python
# TODO: Use the data object and then train the logistic regression model. 
# 1. Change the code below to build your model called LogReg
# 2. Fit the model to the train data

#LogReg = None

...
```

<!-- BEGIN QUESTION -->

#### 3-c. Print out prediction probability and prediction labels from the above model (from the sklearn library) using test data. [5 pts]
Explain 1) why there are two columns in the prediction probability output, and 2) how you can manually optain prediction label from the predictio probability output. 

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

# Part C. Understanding classification performance metrics

### Q4. ROC curve  [10 pts] 
In the next cell, compute the ROC curve and the area under the curve and plot the ROC curve. Your ROC curve plot also should display area under the curve.       
Hint: Use relevant functions in sklearn.metrics. Feel free to refer to the sklearn documentation's examples.

<!-- grade_id: cell-5002bb7d0edf9fc7, points: 0, grade: false, solution: true -->
```python
# TODO: compute the area under the curve and plot ROC curve
# Plot the ROC curve ( True positive rate v/s False positive rate) and indicate the AUC on the plot


...
```

<!-- END QUESTION -->

### Q5. Precision and Recall [10 pts]
Here we will use the trained model coefficients and generate the `classification probabilities` using the `gen_logistic` function we built. The goal of this section is to make you understand how logistic regression classifies data points during and after training. Using the predictions from the generated probabilities, you will compute the precision and recall metrics (defined below).

#### Q5a. Complete the code below to implement precision function. [5 pts]
Assume positive label is 1 and negative label is 0

<!-- grade_id: cell-2332ba460f281cde, points: 0, grade: false, solution: true -->
```python
def calculate_precision(y_true, y_pred):
    '''
    This function accepts the labels and the predictions, then
    calculates precision for a binary classifier.
    
    Args
        y_true: np.ndarray
        y_pred: np.ndarray
    
    Returns precision as a floating point number between 0.0 and 1.0
    '''
    ...

```

```python
grader.check("q5a")
```

#### Q5b. Complete the code below to implement recall function. [5 pts]
 Assume positive label is 1 and negative label is 0

```python
def calculate_recall(y_true, y_pred):
    '''
    This function accepts the labels and the predictions, then
    calculates recall for a binary classifier.
    
    Args
        y_true: np.ndarray
        y_pred: np.ndarray
        
    
    Returns precision as a floating point number between 0.0 and 1.0
    '''
    
    ...
```

```python
grader.check("q5b")
```

<!-- BEGIN QUESTION -->

### Q6. Putting things together [10 pts]

In the next cell you will generate the predictions for the test data `data.x_test` and compute prediction and recall metrics by calling the functions you built above. 
STEP1. Get weight and bias from your fitted model (sklearn model)     
STEP2. Plug weight and bias and test data into your `gen_logistic` function to get prediction probability.    
STEP3. From the predicion probability output from STEP2, calculate prediction label (y_pred).    


<!-- grade_id: cell-176f3759a371e25b, points: 0, grade: false, solution: true -->
```python
# TO-DO : Generate predicted y values using coefficients of the fit logistic regression model for data.x_test
# Then compute and print the precision and recall metrics 

...

# Checking your results. Do not modify codes below.
print(y_pred.shape)
precision = calculate_precision(data.y_test, y_pred)
recall = calculate_recall(data.y_test, y_pred)



print('Model Precision : %0.2f' % precision)
print('Model Recall : %0.2f' % recall)

```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q2b (points: None)


Tests for q2c (points: None)


Tests for q2d (points: None)


Tests for q3a (points: None)


Tests for q3b (points: None)


Tests for q5a (points: None)
```
>>> ut_true = np.array([1.0, 1.0, 0.0, 0.0])
>>> ut_pred = np.array([1.0, 1.0, 1.0, 1.0])
>>> prec = calculate_precision(ut_true, ut_pred)
>>> 
>>> assert prec == 0.5, "Checks the prec value returned from your calculate_precision function."
```


Tests for q5b (points: None)
```
>>> ut_true = np.array([1.0, 1.0, 0.0, 0.0])
>>> ut_pred = np.array([1.0, 1.0, 1.0, 1.0])
>>> recall = calculate_recall(ut_true, ut_pred)
>>> 
>>> assert recall == 1.0, "Checks the recall value returned from your calculate_recall function."
```


Tests for qA (points: None)
