```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

<a id='verytop'></a>

# Homework 1
## Intro to Jupyter Notebooks and Python
## Submit on Gradescope


### Detailed Submission Instructions Are Provided at the end of this Notebook


## Assignment Scope:


1. Navigate Jupyter notebooks (like this one);
2. Write and evaluate some basic *expressions* in Python, the computer language of the course;
3. Practice with basic  *functions*
4. Practice writing math expressions with LaTeX


This hw covers parts of [Chapter 3 ](http://www.inferentialthinking.com/chapters/03/programming-in-python.html) and [Chapter 8](https://inferentialthinking.com/chapters/08/Functions_and_Tables.html) of the online textbook: Computational and Inferential Thinking: The Foundations of Data Science


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Please see the ***Course Syllabus for the Collaboration Policy***. 

The following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by an Instructor or on a quiz/exam.   

We are here to help!  Visit HW Hours and/or post questions on Piazza!

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Moodle please list those references here:

**External references**:  *list any websites you referenced*

## Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1 | 5 | manual
Question 2 | 5 | autograded
Question 3 | 5 | autograded
Question 4 | 5 | autograded
Question 5 | 5 | autograded
Question 6 | 5 | autograded
Question 7 | 5 | autograded
Question 8 | 5 | autograded
Question 9 | 2.5 | autograded
Question 10 | 2.5 | autograded
Question 11 |5 | manual
<a id='top'></a>Total | 50

<a id='top'></a>
---
**Shortcuts:**  [Problem 1](#p1) | [Problem 2](#p2) | [Problem 3](#p3) | [Problem 4](#p4) [Problem 5](#p5) | [Problem 6](#p6) | [Problem 7](#p7) | [Problem 8](#p8)| [Problem 9](#p9) | [Problem 10](#p10) | [Problem 11](#p11)

---

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# 1. Jupyter Notebooks
This webpage is called a Jupyter Notebook.  Jupyter Notebook is an open source web application that you can use to create and share documents that contain live code, equations, visualizations, and text.  
(Fun fact: The name, Jupyter, is a loose acronym meaning Julia, Python, and R, although now it supports many other languages).     
We’ll be using Python in this class. 



### Jupyter Shortcuts ###

Here are some useful Jupyter notebook keyboard shortcuts.  To learn more keyboard shortcuts, go to **Help -> Keyboard Shortcuts** in the menu above. 

Here are a few we like:
1. `ctrl`+`return` : *Evaluate the current cell*
1. `shift`+`return`: *Evaluate the current cell and move to the next*
1. `esc` : *command mode* (may need to press before using any of the commands below)
1. `esc`+`a` : *create a cell above*
1. `esc`+`b` : *create a cell below*
1. `esc`+`dd` : *delete a cell*
1. `esc`+`m` : *convert a cell to markdown*
1. `esc`+`y` : *convert a cell to code*

## 1.1. Text cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings. 

After you edit a text cell, click the "run cell" button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. 

**Practice 1.1.1.** This paragraph is in its own text cell.  Try editing it so that this sentence is the last sentence in the paragraph, and then click the "run cell" ▶| button or hold down `shift` + `return`.  This sentence, for example, should be deleted.  So should this one.

**Practice 1.1.2**:   Double click on this cell to see the Markdown "behind the scenes"</span>:
 
This is type

<span style='color:Green'> This is GREEN type </span>

<span style='background :yellow' > This is highlighted type </span>

### This is big type
## This is bigger type
# This is even bigger type


*Italic*  or _Italic_ 

**Bold** 

***Bold and Italic*** 

~~Scratch Me.~~


[Back to top](#top)

<a id='p1'></a>


<!-- BEGIN QUESTION -->

### Question 1 (5 pts): Getting to know you

We'd like to learn about your math background.  Answer both of the questions in the single cell provided below.  (Please do not add any additional cells for your answer).  

a).  How do you learn math best?

b).  What was the last math class you took?  Where and when did you take it and how did it go?

After you have typed your answer in the cell run the cell (hold down shift + return).  


_Type your answer here, replacing this text._

<!-- END QUESTION -->

[Back to top](#top)

<a id='p2'></a>

### Question 2 (5 pts):  


**STEP 1**  If you haven't already, scroll up to the top of this notebook and run the first cell that says Initialize Otter:   [Back to top](#verytop) 

**STEP 2** 

In the Question 2 thread on Piazza, there will be a 5-digit integer code displayed. Assign `q2_int` to that integer below (i.e. replace the ellipses (`...`) with the 5-digit code). **After you have typed your answer in the cell run the cell (hold down shift + return)**.  


Running the `grader.check` cell directly after it will test whether you have assigned `q2_int` to an integer (Note that this test won't determine whether you have assigned it the correct integer code - that will be done via a hidden test that is run once the assignment has been submitted to Gradescope).  


```python
q2_int = ...
```

```python
grader.check("q2")
```

## 1.2. Code cells
Other cells contain code in the Python 3 language. Running a code cell will execute all of the code it contains.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶| or hold down `shift` + `return`.

**Practice 1.2.1** Try running the cell below:

```python
print("Hello, World!")
```

And this one:

```python
print("\N{WAVING HAND SIGN}, \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

**Expressions:**  The fundamental building block of Python code is an expression. Cells can contain multiple lines with multiple expressions. When you run a cell, the lines of code are executed in the order in which they appear. Every `print` expression prints a line. 

Run the next cell and notice the order of the output.

```python
print("First this line is printed,")
print("and then this one.")
```

## 1.3. Writing Jupyter notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar.  It'll start out as a text cell.  You can change it to a code cell by clicking inside it so it's highlighted, clicking the drop-down box next to the restart (⟳) button in the menu bar, and choosing "Code".

**Practice 1.3.1.** Add a code cell below this one.  Write code in it that prints out:
   

    First this line,
    then the whole 🌏,
    and then this one.


Run your cell to verify that it works.

## 1.4. Errors
Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:
1. The rules are *simple*.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.
2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

Errors are okay; even experienced programmers make many errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

```python
print("This line is missing something."
```

**Note:** In the toolbar, there is the option to click `Cell > Run All`, which will run all the code cells in this notebook in order. However, the notebook stops running code cells if it hits an error, like the one in the cell above.

You should see something like this (minus our annotations):

<img src="error.jpg"/>

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you have created an illegal structure.  The message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it.  (Of course, if you're frustrated, ask a neighbor or a staff member for help.) 


**Practice 1.4.1** Try to fix the code above so that you can run the cell and see the intended message instead of an error.

## 1.5. The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

Next to every code cell, you'll see some text that says `[ ]`. Before you run the cell, you'll see `[ ]`. When the cell is running, you'll see `[*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `[1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection.  Or if you find that when you try to run the top cell in this notebook it outputs an error it's an issue with the kernel.  If this happens, try the following steps:

1. At the top of your screen, click **Kernel**, then **Interrupt**.
2. If that doesn't help, click **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.

# 2. Number Types

Quantitative information arises everywhere in data science. In addition to representing commands to print out lines, expressions can represent numbers and methods of combining numbers.


Python has two real number types:

`int`: 		an integer of any size

`float`:	a number with an optional fractional part

An int never has a decimal point; a float always does

A float might be printed using scientific notation




The expression `3.2500` evaluates to the floating point 3.25. (Run the cell and see.)

```python
3.2500
```

Notice that we didn't have to `print`. When you run a notebook cell, if the last line has a value, then Jupyter helpfully prints out that value for you. However, it won't print out prior lines automatically.

```python
print(2)
3
4
```

Above, you should see that 4 is the value of the last expression, 2 is printed, but 3 is lost forever because it was neither printed nor last.

You don't want to print everything all the time anyway.  But if you feel sorry for 3, change the cell above to print it.

## 2.1 Limitations of float values:

Be aware that float values have the following limitations, which can lead to issues if you're not aware:

- They have limited precision of 15-16 decimal places

- After arithmetic, the final few decimal places can be wrong

- They have limited size (but the limit is huge)


**Practice 2.2.1**:  Run the cells below to observe some of the issues that arise when working with float values

```python
382764238764.238742034876503865238723985 - 382764238764.238742
```

```python
print(0.1 + 0.2)
```

```python
0.1 + 0.2 == 0.3
```

# 3. Names and Assignment Statements
In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large mammal with brown fur and sharp teeth!"  Instead, we just say, "Bear!"

In Python, we do this with *assignment statements*. An assignment statement has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
ten = 3 * 2 + 4
```

When you run that cell, Python first computes the value of the expression on the right-hand side, `3 * 2 + 4`, which is the number 10.  Then it assigns that value to the name `ten`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `ten`:

```python
ten
```

The statement `ten = 3 * 2 + 4` is not asserting that `ten` is already equal to `3 * 2 + 4`, as we might expect by analogy with math notation.  Rather, that line of code changes what `ten` means; it now refers to the value 10, whereas before it meant nothing at all.

If the designers of Python had been ruthlessly pedantic, they might have made us write

    define the name ten to hereafter have the value of 3 * 2 + 4 

instead.  You will probably appreciate the brevity of "`=`"!  But keep in mind that this is the real meaning.

**Practice 3.1.** Run the following cell which uses a variable name `eleven` that hasn't been assigned to anything. You'll see an error!

```python
eleven + 8
```

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output. 

```python
close_to_pi = 355/113
close_to_pi
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.

```python
semimonthly_salary = 842.5
monthly_salary = 2 * semimonthly_salary
number_of_months_in_a_year = 12
yearly_salary = number_of_months_in_a_year * monthly_salary
yearly_salary
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).  And names can't contain spaces, since spaces are used to separate pieces of code from each other.

Other than those rules, what you name something doesn't matter *to Python*.  For example, this cell does the same thing as the above cell, except everything has a different name:

```python
a = 842.5
b = 2 * a
c = 12
d = c * b
d
```

**However**, names are very important for making your code *readable* to yourself and others.  The cell above is shorter, but it's totally useless without an explanation of what it does.

## 3.1. Checking Your Code

Now that you know how to name things, you can start using the built-in *tests* to check whether your work is correct. Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question.  
**Some questions will also contain hidden tests that won't be run until you submit to Gradescope.**   Please don't change the contents of the test cells. 

Run the cell below.  It will import the module NumPy which will be used in some of the tests below.  We will discuss more about NumPy and importing modules in a later section of this assignment.

```python
import numpy as np
```

Go ahead and attempt Question 3 below. Running the cell directly after it will test whether you have assigned `seconds_in_a_decade` correctly in Question 3. If you haven't, this test will tell you the correct answer. Resist the urge to just copy it, and instead try to adjust your expression. (Sometimes the tests will give hints about what went wrong...)

<a id='p3'></a>

### Question 3 (5 points)
Assign the name `seconds_in_a_decade` to the number of seconds between midnight January 1, 2010 and midnight January 1, 2020. Note that there are two leap years in this span of a decade. A non-leap year has 365 days and a leap year has 366 days.

Replace the ellipses (`...`) with your expression.


```python
# Change the next line 
# so that it computes the number of seconds in a decade 
# and assigns that number the name, seconds_in_a_decade.

seconds_in_a_decade = ...

# We've put this line in this cell 
# so that it will print the value you've given to seconds_in_a_decade when you run it.  
# You don't need to change this.
seconds_in_a_decade
```

```python
grader.check("q3")
```

[Back to top](#top)

<a id='p4'></a>


### Question 4 (5 pts)
When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q4` below to indicate your answer.  Replace the ellipses (`...`) with your answer.  
1. Python is smart and already knows `4 = 2 + 2`.

2. In Python, it's a rule that the `=` sign must have a variable name to its left, and `4` isn't a variable name.

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.


```python
names_q4 = ...

# We've put this line in this cell 
# so that it will print the value you've given to names_q4 when you run it.  
# You don't need to change this.
names_q4

#Note that the grader check below only checks if you have entered a valid input (i.e. an integer between 1 and 4),
#not whether or not you chose the correct input - that will be graded via a hidden test in Gradescope.
```

```python
grader.check("q4")
```

## 3.2. Operators


Many basic arithmetic operations are built into Python.  The textbook section on [Expressions](http://www.inferentialthinking.com/chapters/03/1/Expressions.html) describes all some key arithmetic operators.  The common operator that differs from typical math notation is `**`, which raises one number to the power of the other. So, `2**3` stands for $2^3$ and evaluates to 8. 

**Practice:** Play around with the operators $\%, /, **$, and $//$ to see what they do:

```python

```

The order of operations is the same as what you learned in elementary school, and Python also has parentheses.  For example, compare the outputs of the cells below. The second cell uses parentheses for a happy new school year!

```python
7+6*5-6*3**2*2**3/4*7
```

```python
7+(6*5-(6*3))**2*((2**3)/4*7)
```

In standard math notation, the first expression is

$$6 + 6 \times 5 - 6 \times 3^2 \times \frac{2^3}{4} \times 7,$$

while the second expression is

$$6 + (6 \times 5 - (6 \times 3))^2 \times \left(\frac{\left(2^3\right)}{4} \times 7\right).$$


[Back to top](#top)

<a id='p5'></a>
    
### Question 5 (5 pts) ###

Write a Python expression in this next cell that's equal to $\displaystyle 5 \times \left(3 \frac{10}{11}\right) - 50 \frac{1}{3} + 2^{0.5 \times 22} - \frac{7}{33} + 6$.  That's five times three and ten elevenths, minus fifty and a third, plus two to the power of half twenty-two, minus seven thirty-thirds plus six.  By "$\displaystyle 3 \frac{10}{11}$" we mean $\displaystyle 3+\frac{10}{11}$, not $\displaystyle 3 \times \frac{10}{11}$.

Replace the ellipses (`...`) with your expression.  Try to use parentheses only when necessary.



```python
q5_answer = ...

# We've put this line in this cell 
# so that it will print the value you've given to q5_answer when you run it.  
# You don't need to change this.
q5_answer
```

```python
grader.check("q5")
```

## 3.3 Comments
You may have noticed these lines in the cell in which you answered Question 1:

    # Change the next line 
    # so that it computes the number of seconds in a decade 
    # and assigns that number the name, seconds_in_a_decade.
    
This is called a *comment*. It doesn't make anything happen in Python; Python ignores anything on a line after a `#`.  Instead, it's there to communicate something about the code to you, the human reader. Comments are extremely useful. 

<img src="http://imgs.xkcd.com/comics/future_self.png">
Source: http://imgs.xkcd.com/comics/future_self.png

# 4. Functions

The most common way to combine or manipulate values in Python is by calling functions. 


Functions are either built-in, or user-defined as demonstrated below. 


### 4.1 Built In Functions:
Python comes with many built-in functions that perform common operations.

Examples of a few built-in functions we will frequently use include:  
`abs(), sum(), min(), max(), len(), set(), sorted(), enumerate(), type(), map(), int() lambda(), round()`


For example, the `abs` function takes a single number as its argument and returns the absolute value of that number. Run the next two cells and see if you understand the output.

```python
abs(5)
```

```python
abs(-5)
```

You can check the type of any object in Python using the built-in `type` function:

```python
type(2)
```

```python
type(2.5)
```

```python
type('hello')
```

```python
type([1,2.5,'hello'])
```

### 4.2 Multiple arguments
Some functions take multiple arguments, separated by commas. For example, the built-in `max` function returns the maximum argument passed to it.

```python
max(2, -3, 4, -5)
```

### 4.3 Viewing Documentation

To output the documentation for a function, use the `help` function.

```python
help(max)
```

You can also use Jupyter to view function documentation inside your notebook. The function must already be defined in the kernel for this to work.

**Practice 4.1.1:** Below, click your mouse anywhere on the `sum` block below and use `Shift` + `Tab` to view the function's documentation. 

```python
max(2, -3, 4, -5)
```

### 4.4. User Designed Functions

Let's write a very simple function that converts a proportion to a percentage by multiplying it by 100.  For example, the value of `to_percentage(.5)` should be the number 50 (no percent sign).

A function definition has a few parts.

##### `def`
It always starts with `def` (short for **def**ine):

    def

##### Name
Next comes the name of the function.  Like other names we've defined, it can't start with a number or contain spaces. Let's call our function `to_percentage`:
    
    def to_percentage

##### Signature
Next comes something called the *signature* of the function.  This tells Python how many arguments your function should have, and what names you'll use to refer to those arguments in the function's code.  A function can have any number of arguments (including 0!). 

`to_percentage` should take one argument, and we'll call that argument `proportion` since it should be a proportion.

    def to_percentage(proportion)
    
If we want our function to take more than one argument, we add a comma between each argument name. Note that if we had zero arguments, we'd still place the parentheses () after that name. 

We put a **colon** after the signature to tell Python that the next indented lines are the body of the function. If you're getting a syntax error after defining a function, check to make sure you remembered the colon!

    def to_percentage(proportion):

##### Documentation
Functions can do complicated things, so you should write an explanation of what your function does.  For small functions, this is less important, but it's a good habit to learn from the start.  Conventionally, Python functions are documented by writing an **indented** triple-quoted string:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""

##### Body
Now we start writing code that runs when the function is called.  This is called the *body* of the function and every line **must be indented with a tab**.  Any lines that are *not* indented and left-aligned with the def statement is considered outside the function. 

Some notes about the body of the function:
- We can write code that we would write anywhere else.  
- We use the arguments defined in the function signature. We can do this because we assume that when we call the function, values are already assigned to those arguments.
- We generally avoid referencing variables defined *outside* the function. If you would like to reference variables outside of the function, pass them through as arguments!


Now, let's give a name to the number we multiply a proportion by to get a percentage:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""
        factor = 100

##### `return`
The special instruction `return` is part of the function's body and tells Python to make the value of the function call equal to whatever comes right after `return`.  We want the value of `to_percentage(.5)` to be the proportion .5 times the factor 100, so we write:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""
        factor = 100
        return proportion * factor
        
`return` only makes sense in the context of a function, and **can never be used outside of a function**. `return` is always the last line of the function because Python stops executing the body of a function once it hits a `return` statement. If a function does not have a return statement, it will not return anything; if you expect a value back from the function, make sure to include a return statement. 

*Note:*  `return` inside a function tells Python what value the function evaluates to. However, there are other functions, like `print`, that have no `return` value. For example, `print` simply prints a certain value out to the console. 

In short, `return` is used when you want to tell the *computer* what the value of some variable is, while `print` is used to tell you, a *human*, its value.

[Back to top](#top)

<a id='p6'></a>

### Question 6 (5 pts) ###

Define `to_percentage` in the cell below. 

Call your function to convert the proportion 0.2 to a percentage.  Name that percentage `twenty_percent`.

Replace EACH ellipses below (`...`)  with the corresponding correct code.  


```python
def ...
    ''' ... '''
    ... = ...
    return ...

twenty_percent = ...
twenty_percent
```

```python
grader.check("q6")
```

Here's something important about functions: the names assigned *within* a function body are only accessible within the function body. Once the function has returned, those names are gone.  So even if you created a variable called `factor` and defined `factor = 100` inside of the body of the `to_percentage` function and then called `to_percentage`, `factor` would not have a value assigned to it outside of the body of `to_percentage`:

**Note**: Below, you should see an error message.

```python
# You should get an error when you run this.  (If you don't, 
# you might have defined factor somewhere above.)
factor
```

[Back to top](#top)

<a id='p7'></a>

### Question 7 (5 pts):

You can also pass in named values as arguments to your function.

Use `to_percentage` again to convert the proportion named `a_proportion` (defined below) to a percentage called `a_percentage`.

*Note:* You don't need to define `to_percentage` again!  Like other named values, functions stick around after you define them.

```python
a_proportion = 2**(0.5) / 2
a_percentage = ...
a_percentage
```

```python
grader.check("q7")
```

[Back to top](#top)

<a id='p8'></a>

### Question 8 (5 pts):

Write a function `square` that takes a number as its argument and returns the square of its argument.

```python
def square(x):
    ...
```

```python
grader.check("q8")
```

### 4.5 Methods 

 
A method is a function associated with a specified object (and, by extension, with the type of that object). That is, a method corresponds to a data-type operation.  Methods are defined within the class from which an object is instantiated. 


We can call (or invoke) a method by using a variable name, followed by the dot operator (.), followed by the method name, followed by a list of arguments delimited by commas and enclosed in parentheses. 

**Practice 4.3.1** As a simple example, Python's built-in int type has a method named bit_length(), so you can determine the number of bits in the binary representation of an int value as follows:


```python
x = 3 ** 100
bits = x.bit_length()
bits
```

## 5. Importing Code

Most programming involves work that is very similar to work that has been done before.  Since writing code is time-consuming, it's good to rely on others' published code when you can.  Rather than copy-pasting, Python allows us to **import modules**. A module is a file with Python code that has defined variables and functions. By importing a module, we are able to use its code in our own notebook.

Python includes many useful modules that are just an `import` away.  


### 5.1:  NumPy module:

𝐍𝐮𝐦𝐏𝐲  (short for "Numeric Python") is an open source Python module for scientific computing. NumPy supports large, multidimensional arrays and matrices. It also supports a large collection of mathematical functions not found in Python's standard math library.

It's custom in data science to import Numpy  with the alias $\texttt{np}$.  We can then ***access any mathematical function in the Numpy library by prepending function names by $\texttt{np}$***

For a more in-depth description of methods and attributes built into Numpy see here:  https://numpy.org/doc/stable/user/quickstart.html   
and  here:  https://numpy.org/doc/stable/reference/routines.math.html

*Hint: If you press `shift+tab` while next to the function call, the documentation for that function will appear.*

In Sec 3.1 above, the line `import numpy as np` imports the numpy module. This statement creates a module and then assigns the name `np` to that module. We are now able to access any variables or functions defined within `np` by typing the name of the module followed by a dot, then followed by the name of the variable or function we want.

    <np>.<name>

```python
radius = 5
area_of_circle = radius**2 * np.pi
area_of_circle
```

[Back to top](#top)

<a id='p9'></a>


### Question 9 (2.5 pts)
The module `numpy` also provides the name `e` for the base of the natural logarithm, which is roughly 2.71. 

Compute $e^{\pi}-\pi$, giving it the name `near_twenty`.

*Remember: You can access `pi` from the `numpy` module as well!*


```python
near_twenty = ...
near_twenty
```

```python
grader.check("q9")
```

### 5.2  Accessing Functions

In the question above, you accessed variables within the `numpy` module. 

**Modules** also define **functions**.  For example, `numpy` provides the name `floor` for the floor function.  Having imported `numpy as np` already, we can write `np.floor(7.5)` to compute the floor of 7.5.  (Note that the floor function returns the largest integer less than or equal to a given number.)


```python
np.floor(7.5)
```

[Back to top](#top)

<a id='p10'></a>

### Question 10 (2.5 pts)

Calculate the square root of `pi` using `sqrt` and `pi` from the `numpy` module.  Give the result the name `sqrt_of_pi`.

```python
sqrt_of_pi = ...
sqrt_of_pi
```

```python
grader.check("q10")
```

### 5.3  Methods vs Functions in NumPy

Many things are implemented in NumPy as both functions and methods:

Examples are `mean`, `min`, `max`, `sum`, `reshape`

For example, there is a `np.sum` function, that adds up all the elements:

```python
arr = np.array([1, 2, 0, 1])
np.sum(arr)
4
```

There is also a `sum()` method of the numpy array object:

```python
type(arr)
```

```python
arr.sum()
```

In these cases choosing the method or the function will usually depend on which one is easier to read.

## 6. Math Questions

### 6.1 LaTex ###
You should use LaTeX to format math in your answers. If you aren't familiar with LaTeX, not to worry. It's not hard to use in a Jupyter notebook. Just place your math in between dollar signs:

\\$ f(x) = 2x \\$ becomes $ f(x) = 2x $.

If you have a longer equation, use double dollar signs to place it on a line by itself:

\\$\\$ \sum_{i=0}^n i^2 \\$\\$ becomes:

$$ \sum_{i=0}^n i^2 $$.

Here is some handy LaTex:

| Output | Latex   |
|:--|:--|
| $$x^{a + b}$$  | `x^{a + b}` |
| $$x_{a + b}$$ | `x_{a + b}` |
| $$\frac{a}{b}$$ | `\frac{a}{b}` |
| $$\sqrt{a + b}$$ | `\sqrt{a + b}` |
| $$\{ \alpha, \beta, \gamma, \pi, \mu, \sigma^2  \}$$ | `\{ \alpha, \beta, \gamma, \pi, \mu, \sigma^2  \}` |
| $$\int_{x=2}^{\infty} \frac{1}{x} \,dx $$ | `\sum_{x=2}^{\infty} \frac{1}{x} \,dx` |
| $$\sum_{x=1}^{100} x$$ | `\sum_{x=1}^{100} x` |
| $$\frac{d}{d x} $$ | `\frac{d}{dx} ` |



[For more about basic LaTeX formatting, you can read this article.](https://www.sharelatex.com/learn/Mathematical_expressions)


<!-- BEGIN QUESTION -->

[Back to top](#top)

<a id='p11'></a>


### Question 11 (5 pts)

Exponential functions and their derivatives will be used in this class when we discuss probability distributions.  

To brush up on Calculus, evaluate the following derivative by hand: $$\frac{d}{dx} \left( e^{x/4} \right)$$

Type your answer in the cell below using Markdown and LaTeX

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 1!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 1 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW1 Autograded assignment in Gradescope.

* **Homework 1 Manually Graded**: Submit your hw01.PDF to the HW1 Manually Graded assignment in Gradescope.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw01.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q10 (points: 2.5)
```
>>> assert np.isclose(sqrt_of_pi, 1.7724538509055159)
```


Tests for q2 (points: [0])
```
>>> # TEST
>>> assert type(q2_int)==int
```


Tests for q3 (points: 5)
```
>>> assert seconds_in_a_decade == 315532800
```


Tests for q4 (points: [0])
```
>>> # TEST
>>> 
>>> assert names_q4 in [1,2,3,4]
```


Tests for q5 (points: 5)
```
>>> assert q5_answer == 2023.0
```


Tests for q6 (points: 5)
```
>>> # Make sure your function has the proper syntax!
>>> assert to_percentage(.35) == 35.0
```
```
>>> # Make sure you are passing in the proper argument!
>>> assert twenty_percent == 20.0
```


Tests for q7 (points: 5)
```
>>> assert np.isclose(a_percentage, 70.71067811865476)
```


Tests for q8 (points: 5)
```
>>> assert square(1) == 1
```
```
>>> assert square(0) == 0
```


Tests for q9 (points: 2.5)
```
>>> assert np.isclose(near_twenty, 19.99909997918947)
```
