# CSPB 3104 Assignment 2:

***
# Instructions

This assignment is to be completed as a python3 notebook.  

The questions  provided  below will ask you to either write code or 
write answers in the form of markdown.

 Markdown syntax guide is here: [click here](https://github.com/adam-p/markdown-here/wiki/Markdown-Cheatsheet)

Using markdown you can typeset formulae using latex.
This way you can write nice readable answers with formulae like thus:

The algorithm runs in time $\Theta\left(n^{2.1\log_2(\log_2( n \log^*(n)))}\right)$, 
where $\log^*(n)$ is the inverse _Ackerman_ function.

__Double click anywhere on this box to find out how your instructor typeset it. Press Shift+Enter to go back.__


***
## Question 1: Setting Up and Solving Recurrences

Consider the python-like pseudocode below

~~~
def div_and_conquer_fun(a):
    # a is an array of size n
    n = length(a)
    if n == 0:
        return 0
    if n == 1: 
        return a[1]
    # 1. Divide into 3 parts
    a1 = a[1 ... n//3]
    a2 = a[n//3+1 ... 2*n//3]
    a3 = a[2*n//3+1 ... n]
    # note // denotes integer division a//b := floor(a/b)
    (b1, b2) = coalesce_arrays_into_two(a1, a2, a3)
    # note b1, b2 are arrays of size n//4 each.
    c1 = div_and_conquer_fun(b1)
    c2 = div_and_conquer_fun(b2)
    return c1 + c2 // Theta (1) time
~~~

1. The algorithm first divides an array of size n into 3 roughly equal parts.
2. Next, it uses the function `coalesce_arrays_into_two(a1, a2, a3)` that runs in $\Theta(n)$ time, returning two arrays `b1` and `b2` of size $\frac{n}{4}$ each.
3. The function is then recursively called on `b1` and `b2`.
4. Finally, the result is summed up and returned.

Write down a recurrence relation for the running time of the divide and conquer function above. Use master method to solve the recurrence: write down which case of the master method and the result.


<!-- grade_id: cell-0025f295bb34e01b, points: 20, grade: true, solution: true -->
To solve this problem, the first step is finding a recurrence relation for our algorithm. In general, the formula for a recurrence (excluding the base case) is

\begin{equation} \tag{1}
T(n) = aT(n / b) + f_{\text{div}} + f_{\text{comb}}
\end{equation}

where $f_{\text{div}}$ is the time complexity that it takes for a function to divide and $f_{\text{comb}}$ is the time complexity that it takes for a function to combine. We now need to find the values for $a$ and $b$ and the time complexity for this pseudo code to divide and then combine.

From the problem statement we know that this algorithm is dividing a single array into three roughly equal parts. It then calls a function where it combines three sub-arrays into two arrays of $n / 4$ size. The algorithm then makes two recursive calls. To fill in the values for $a$ and $b$ in our aforementioned recurrence formula, we need to think about what is happening for each recurrence operation.

For each recurrence

- $a$ - The number of recursive calls in this algorithm is 2. So therefore $a = 2$.
- $b$ - How the current problem is being reduced in each step. We are told that `coalesce_arrays_into_two` is returning arrays of $n / 4$. Therefore $b = 4$.
- $f_{\text{div}}$ - The division step essentially runs at constant complexity, namely $\Theta(1)$.
- $f_{\text{comb}}$ - The combination step (`coalesce_arrays_into_two`) is combining arrays in $\Theta(n)$.

With this information we can then say $T(n)$ is

\begin{equation} \tag{2}
T(n) = 2T(n / 4) + \Theta(1) + \Theta(n).
\end{equation}

We now move on to solving this recurrence. In general, when proceeding to use the Master's Theorem, we need to compare the value of $\log_{b}{(a)}$ to that of a constant $c$. Earlier, we stated that the formula for a recurrence relation was $T(n) = aT(n / b) + f_{\text{div}} + f_{\text{comb}}$. In reality, we only really care about the dominant term from $f_{\text{div}}$ and $f_{\text{comb}}$ in (1). We can re-use this previous information to then say

\begin{equation} \tag{3}
T(n) = aT(n / b) + \Theta(n^{c})
\end{equation}

for our specific case. This then means that the recurrence relation for our pseudo code is going to be

\begin{equation} \tag{4}
T(n) = 2 T(n / 4) + \Theta(n^{1}).
\end{equation}

We now are ready to solve this recurrence with the Master's Theorem. The Master's Theorem states that for a given value of $\epsilon = \log_{b}{(a)}$, the runtime complexity of the entire algorithm is

- Case 1: $\epsilon > c \hspace{10pt} \therefore \hspace{10pt} \Theta(n^{\epsilon})$
- Case 2: $\epsilon = c \hspace{10pt} \therefore \hspace{10pt} \Theta(n^{\epsilon}\log{(n)})$
- Case 3: $\epsilon < c \hspace{10pt} \therefore \hspace{10pt} \Theta(n^{c})$

In our case, $\epsilon$ is calculated to be

\begin{equation} \tag{5}
\epsilon = \log_{b}{(a)} = \log_{4}{(2)} = 0.5
\end{equation}

Using a quick comparison we can see that $\epsilon = 0.5 < c = 1$, so therefore we need to use Case 3 from the Master's Theorem. Therefore our solution for this pseudo code in terms of runtime complexity is going to be

\begin{equation} \tag{6}
\color{blue}{\Theta(n)}.
\end{equation}

***
## Question 2(a): Counting Dominances
Suppose you are given two sorted arrays $a$ and $b$ of the sizes $m$ and $n$, respectively. A "dominance" of $a$ over $b$ is a pair of indices $(i,j)$ such that $a[i] > b[j]$.  Note that $i$ is an index of array $a$ and $j$ must be an index of array $b$.


Write a __brute force__ algorithm that counts the number of dominances of $a$ over $b$ that runs in $\Theta(n^2)$ time.

<!-- grade_id: cell-57660028cdf2b31f, points: 0, grade: false, solution: true -->
```python
#Answer 2(a):
""" count_dominances_brute_force - Counts how many elements of a are greater than that of b
    Input:
        a - Array of values
        b - Array of values
    Algorithm:
        * Create a counter for the number of dominances
        * Iterate over each element of a
            * Iterate over element of b
                * If the element in a is greater than that of the element in b, increment the counter by 1
        * Return the counter
    Output:
        dom - Counter for the number of dominances between arrays a and b
"""
def count_dominances_brute_force(a, b):
    # YOUR CODE HERE
    # Dominance counter
    dom = 0
    # Iterate over a
    for i in range(len(a)):
        # Iterate over b
        for j in range(len(b)):
            # Increment counter is a[i] > b[j]
            if (a[i] > b[j]):
                dom += 1
    # Return counter
    return dom
```

## Question 2(b): Counting Dominances
However, the brute force algorithm is suboptimal. Design a $\Theta(n)$ algorithm to count the number of dominances. Do this by modifying the merge algorithm we studied as part of merge sort. Instead of merging the two sorted arrays, count the number of dominances.

<!-- grade_id: cell-a6e9bad601c5a14f, points: 0, grade: false, solution: true -->
```python
#Answer 2(b):
""" count_dominances - Counts dominances of two arrays at O(n)
    Input:
        a - Array of values
        b - Array of values
    Algorithm:
        * Create counters for dominances and indices of arrays
        * Traverse a and b
            * Check if a dominance occurs in the data set
                * If there is a dominance, increment the counter, advance an element in b
            * If it doesn't occur
                * Advance an element in a
    Output:
        dom - Number of dominances between list a and b
"""
def count_dominances(a, b):
    # YOUR CODE HERE
    # Counters
    dom = 0
    i = 0
    j = 0
    # Traverse a and b
    while (i < len(a) and j < len(b)):
        # Dominance
        if (a[i] > b[j]):
            # Increment dominance
            dom += (len(a) - i)
            # Advance index of b
            j += 1
        # No dominance
        else:
            i += 1
    return dom
```

## Question 2(c): Counting Dominances
Explain the logic behind your solution to Question 2(b) briefly (5 lines).

<!-- grade_id: cell-efd6f00eccd2e786, points: 15, grade: true, solution: true -->
The above algorithm is of the time complexity $\mathcal{O}(n)$. This is mainly due to the fact that the each element of the arrays are only traversed once. The above while loop will run until we have reached the end of a or the end of b. Inside the while loop, we check to first see if the current element of a is greater than the current element of b. If it is, then we increment the dominance count based upon the length of a and the current index index of a, and then advance the index of b. If a dominance does not occur, then we simply advance to the next element in a and check the current value (if not at the end of the list).

***
## Question 3(a): Finding a Fixed Point. 
A fixed point of an array $A$, if it exists, is an index $i$ such that $A[i] = i$.
Given a _sorted_ array $A$ of _distinct_ __integers__, return the index of the fixed point if one exists, or otherwise, return `-1` to signal that no fixed point exists. Your algorithm must be as efficient as possible.

<!-- grade_id: cell-6a58c57b10f4d43e, points: 0, grade: false, solution: true -->
```python
#Answer 3(a)
""" find_fixed_point - Finds the fixed point of an array
    Input:
        a - Array of sorted values
    Algorithm:
        * Use a binary search to find the fixed point of an array
            * Create trackers for left and right ends of searched array
            * Traverse the 'split' arrays
                * Calculate the midpoint based off the current left and right indices
                * If the current value in a[mid] is the midpoint (index)
                    * Return the midpoint value
                * If the current value in a[mid] is less than the midpoint
                    * Move the left end of where we are searching to one past the midpoint
                * If the current value in a[mid] is greater than the midpoint
                    * Move the right end of where we are searching to one before the midpoint
            * If no fixed point is round, return -1
    Output:
        This function returns the fixed point of an array if it is found, otherwise it returns -1
"""
def find_fixed_point(a):
    # Left and right starting
    left = 0
    right = len(a) - 1
    # Traverse the array
    while left <= right:
        # Calculate midpoint
        mid = (left + right) // 2
        # Mid point is index
        if a[mid] == mid:
            return mid
        # Point is less than mid point, move left end one past midpoint
        elif a[mid] < mid:
            left = mid + 1
        # Point is greater than mid point, move right end one before midpoint
        else:
            right = mid - 1
    # No mid point found
    return -1
```

## Question 3(b): 
Explain your solution to the problem briefly and derive its running time complexity.

<!-- grade_id: cell-15fd886bc11e6fd3, points: 15, grade: true, solution: true -->
For this problem, I chose to use a binary search algorithm as I remember that these algorithms are extremely efficient from what I learned in Data Structures. This algorithm checks if the midpoint of where we are searching in the current section of the array is equal to the value at the current index of the midpoint. It utilizes the common boiler plate code that is found in binary search algorithms. It correctly increments or decrements the left and right ends of the arrays so that it constantly splits the current array in half for where it is searching until the condition is found, or the left end is greater than the right end.

To derive the running time complexity of this algorithm, we need to determine how many steps this algorithm is going to take until it is finished running. If we consider our array to be of length $n$, and the number of steps to be $m$, then the we have the resulting equation

$$
\frac{n}{2^{m}} = 1
$$

where the above is looking at how many times we can divide $n$ by two until it reaches a single element. We can rearrange the above so that we can solve for $m$ and we have

\begin{align*}
n & = 2^{m} & \text{(Rearranging prior expression)} \\
\log_{2}{(n)} & = \log_{2}(2^{m}) & \text{(Applying \texttt{Log} base 2 to each side)} \\
\log_{2}{(n)} & = m & \text{(Simplification)}.
\end{align*}

This means that the time complexity of our algorithm is going to be

$$
\color{blue}{\mathcal{O}(\log{(n)})}.
$$

## Question 3(c): Finding a Fixed Point Again. 

Given a _sorted_ array $A$ of _distinct_ __natural numbers__, return the index of the fixed point if one exists, or otherwise, return `-1` to signal that no fixed point exists. Your algorithm must be as efficient as possible.

<!-- grade_id: cell-03b15d6700803f4e, points: 0, grade: false, solution: true -->
```python
#Answer 3(c)
""" find_fixed_point_natural - Finds the fixed point of an array for distinct natural numbers
    Input:
        a - Array of sorted natural values
    Algorithm:
        * Check if the first value is a zero
            * If it is, then return that value
        * If it is not a zero
            * No fixed point can exist, therefore return -1
    Output:
        This function returns the first fixed point in an array of natural numbers if the first element is a 0
"""
def find_fixed_point_natural(a):
    # First entry is a zero, fixed point exists
    if (a[0] == 0):
        return 0
    # First entry is not zero, no fixed point can exist
    else:
        return -1
```

## Question 3(d)

Explain your solution to the problem briefly and derive its running time complexity. (*Hint*: Your algorithm should be quicker than your algorithm for part (a).)

<!-- grade_id: cell-f933018b82a90554, points: 10, grade: true, solution: true -->
For this problem, considering that we are using distinct (all different) natural numbers (integers starting from zero), we just need to check if the first element is a zero. If the first element is a zero, then a fixed point does indeed exist and that fixed point is 0. If the first element is not a zero, then we know that since this array is sorted as well that a fixed point cannot exist. So in this case we return -1.

The time complexity of this algorithm is 

$$
\color{blue}{\mathcal{O}(1)}.
$$

This is because there are no loops involved in this algorithm and it will only execute the if / else statement once. So therefore the algorithm executes in constant time. No matter the size of the array $\texttt{a}$, this algorithm will always run in constant time.

## Testing your solutions -- Do not edit code beyond this point

<!-- grade_id: cell-07261dba53118757, points: 0, grade: true, solution: false -->
```python
# This code runs 5 test cases on your two algorithms
def test_count_dominances(func):
    a1 = [ 5, 7, 10]
    b1 = [ 1, 2,  3] 
    n1 = 9

    a2 = [ 6, 10, 15, 21]
    b2 = [ 4, 19, 25, 32]
    n2 = 5
    
    a3 = [ 6, 10, 15, 21]
    b3 = []
    n3 = 0
    
    a4 = [ 1, 3, 5, 7, 9, 11, 13]
    b4 = [ 2, 4, 6, 8, 10]
    n4 = 20
    
    a5 = [1, 3, 5, 6, 7, 9, 11, 13]
    b5 = [2, 4, 6, 6, 6, 8, 10]
    n5 = 30
    
    problems = [(a1, b1, n1), (a2, b2, n2), (a3, b3, n3), (a4, b4, n4), (a5, b5, n5)]
    num_passed = 0
    for (a, b, n) in problems:
        res = func(a, b)
        if res == n:
            num_passed = num_passed + 1
        else: 
            print('FAILED: a = ', a, 'b = ', b, ' expected = ', n, 'your code = ', res)
    print('--- Done ---')
    print ('Num tests = ', len(problems))
    print ('Num passed = ', num_passed)
```

<!-- grade_id: cell-5fec67453127a249, points: 10, grade: true, solution: false -->
```python
print('Testing brute force:')
test_count_dominances(count_dominances_brute_force)
```

Output:
```
Testing brute force:
--- Done ---
Num tests =  5
Num passed =  5
```

<!-- grade_id: cell-a4ebce53d021bc00, points: 10, grade: true, solution: false -->
```python
print('Testing modified merge algorithm:')
test_count_dominances(count_dominances)
```

Output:
```
Testing modified merge algorithm:
--- Done ---
Num tests =  5
Num passed =  5
```

```python
from random import sample
def compare_brute_force_vs_fast():
    a = sorted( sample (range(60), 20) )
    b = sorted( sample (range(60), 20) )
    n1 = count_dominances_brute_force(a, b)
    n2 = count_dominances(a, b)
    if n1 != n2:
        print('Disparity observed between two algorithms:', a, b, n1, n2)
        return False
    return True
    
print('Comparing the two implementations.')
num_passed = 0
total = 100
for i in range(total):
    if compare_brute_force_vs_fast():
        num_passed = num_passed + 1
print(' -- all tests done -- ')
print(' passed = ', num_passed, ' out of ', total)
```

Output:
```
Comparing the two implementations.
 -- all tests done -- 
 passed =  100  out of  100
```

```python
print(find_fixed_point([-10, -5, -2, 2, 3, 5, 7, 10, 15, 25, 35, 78, 129]))
```

Output:
```
5
```

```python
def find_fixed_point_very_naive(a):
    n = len(a)
    for i in range(0, n):
        if a[i] == i:
            return i
    return -1
```

<!-- grade_id: cell-84c230b8d60d7a54, points: 10, grade: true, solution: false -->
```python
def test_find_fixed_point_code(n_tests, test_size):
    n_passed = 0
    for i in range(0, n_tests):
        a = sorted( sample( range(-10 * n_tests,  10 * n_tests ), test_size))
        j = find_fixed_point(a)
        if j >= 0 and a[j] != j:
            print(' Code failed for input: ', a, 'returned : ', j, 'expected:', find_fixed_point_very_naive(a))
        elif j < 0: 
            assert j == -1, 'Your code returns an illegal negative number: have you implemented it yet?'
            k = find_fixed_point_very_naive(a)
            if k >= 0:
                print('Code failed for input', a)
                print('Your code failed to find a fixed point')
                print('However, for j = ', k, 'a[j] =', a[k])
            else: 
                n_passed = n_passed + 1
        else: 
            n_passed = n_passed + 1
    return n_passed

n_tests = 10000
n_passed = test_find_fixed_point_code(10000, 10)
print(' num tests  = ', n_tests)
print(' num passed = ', n_passed)
```

Output:
```
num tests  =  10000
 num passed =  10000
```

```python
print('Test: expected answer = 5, your answer = ', find_fixed_point([-10, -5, -2, 2, 3, 5, 7, 10, 15, 25, 35, 78, 129])) 
```

Output:
```
Test: expected answer = 5, your answer =  5
```

<!-- grade_id: cell-1a18d5990fb957a8, points: 10, grade: true, solution: false -->
```python
def test_find_fixed_point_natural_code(n_tests, test_size):
    n_passed = 0
    for i in range(0, n_tests):
        a = sorted( sample( range(0,  10 * n_tests ), test_size))
        j = find_fixed_point_natural(a)
        if j >= 0 and a[j] != j:
            print(' Code failed for input: ', a, 'returned : ', j, 'expected:', find_fixed_point_very_naive(a))
        elif j < 0: 
            assert j == -1, 'Your code returns an illegal negative number: have you implemented it yet?'
            k = find_fixed_point_very_naive(a)
            if k >= 0:
                print('Code failed for input', a)
                print('Your code failed to find a fixed point')
                print('However, for j = ', k, 'a[j] =', a[k])
            else: 
                n_passed = n_passed + 1
        else: 
            n_passed = n_passed + 1
    return n_passed

n_tests = 10000
n_passed = test_find_fixed_point_natural_code(10000, 10)
print(' num tests  = ', n_tests)
print(' num passed = ', n_passed)
```

Output:
```
num tests  =  10000
 num passed =  10000
```

```python
print('Test: expected answer = 0, your answer = ', find_fixed_point_natural([0,1, 2, 3, 5, 7, 10, 15, 25, 35, 78, 129]))
```

Output:
```
Test: expected answer = 0, your answer =  0
```