```scala

```

# CSPB 3155: Assignment 3 

Topics covered: ASTs, pattern matching, higher order functions.

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (15 points)

Adding a switch statement for the `mython` (var) language developed in class. 

The syntax for the switch statement is given in terms of the following additional rules that
extend the mython implementation found in the class notes.

$$\begin{array}{rcl}
\mathbf{Statement} & \rightarrow & \text{Switch}(\mathbf{Expr},  \mathbf{Case}* ) \\[5pt]
\mathbf{Case} & \rightarrow & \text{SwitchCase}(\mathbf{Expr}, \mathbf{Statement}* ) \\[5pt]
& | & \text{DefaultCase}(\mathbf{Statement}*) \\
\end{array}$$

The box below contains the scala language definitions of `mython` so far, with space for you to 
concretely implement the abstract grammar above.

Augment the definitions by adding the extra definitions to support switch statements. Please use
the same constructor names as in the grammar: `Switch` for the switch statement and 
`SwitchCase` and `DefaultCase` for the two rules corresponding to the __Case__ specifications.

For your convenience, the definitions used in the class are given. Add your solution at the end.


<!-- grade_id: cell-a3a78387d2a93841, points: 0, grade: false, solution: true -->
```scala
sealed trait Expr
case class Const(d: Double) extends Expr 
// 1. We cheated a bit and allowed all floating point numbers
// Also, this deviates from the grammar
case class Ident(s: String) extends Expr
// 2. We allow any string to be an identifier for now instead of the regular expression shown in the grammar.
case class Plus(e1: Expr, e2: Expr ) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Mult(e1: Expr, e2: Expr) extends Expr
case class Div(e1: Expr, e2: Expr) extends Expr
case class Negate(e: Expr) extends Expr
case class Log(e: Expr) extends Expr
case class Exp(e: Expr) extends Expr
case class Sine(e: Expr) extends Expr
case class Cosine(e: Expr) extends Expr

sealed trait CondExpr
case object ConstTrue extends CondExpr
case object ConstFalse extends CondExpr
case class Geq(e1: Expr, e2: Expr) extends CondExpr
case class Leq(e1: Expr, e2: Expr) extends CondExpr
case class Eq(e1: Expr, e2: Expr) extends CondExpr
case class And(c1: CondExpr, c2: CondExpr) extends CondExpr
case class Or(c1: CondExpr, c2: CondExpr) extends CondExpr
case class Not(c: CondExpr) extends CondExpr

sealed trait Declaration
sealed trait Statement
sealed trait MythonProgram

case class Program(decls: List[Declaration], stmts: List[Statement], returnAtEnd: Expr) extends MythonProgram // We stripped the ReturnStmt tag since it is redundant
case class VarDecl(identifier: String, rhsExpr: Expr) extends Declaration
case class AssignStmt(identifier: String, rhsExpr: Expr) extends Statement
case class WhileStmt(cond: CondExpr, stmts: List[Statement]) extends Statement
case class IfThenElseStmt(cond: CondExpr, stmtsThen: List[Statement], stmtsElse: List[Statement]) extends Statement
case class ReturnStmt(retExpr: Expr) extends Statement

??? // YOUR CODE HERE
```

<!-- grade_id: cell-b4433298fac96b06, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Plus(Ident("y"), Plus( Ident("z"), Exp( Minus(Ident("x"), Ident("y")) ) )) // y + z + exp(x - y)
val stmt1 = AssignStmt("x", e1) // x = y + z + exp(x - y)
// Let us make up the parts of the while statement
val cond1 = Leq(Ident("y"), Const(15.0f)) // y <= 15
// The assignment inside the while
val stmt2 = AssignStmt("y", Minus(Ident("y"), Ident("x"))) // y = y - x
// The condition for the if statement
val cond2 = Leq(Ident("x"), Const(0.0f)) // x <= 0.0
val stmt3 = AssignStmt("x", Minus(Const(0.0f), Ident("x"))) // x= 0 - x
val ifStmt = IfThenElseStmt(cond2, List(stmt3), List()) // if ( x <= 0) x = -x else nothing
val whileStmt = WhileStmt(cond1, List(stmt2, ifStmt)) // while loop 
val allstmts = List(stmt1, whileStmt)
val retExpr = Minus(Ident("y"), Ident("x")) // y - x to be returned

val c1 = SwitchCase(Const(2.0), List(stmt1, stmt3))
val c3 = SwitchCase(Ident("y"), List(stmt1, stmt2, whileStmt))

passed(5)
//END TEST
```

<!-- grade_id: cell-590403a05cbd399c, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Plus(Ident("y"), Plus( Ident("z"), Exp( Minus(Ident("x"), Ident("y")) ) )) // y + z + exp(x - y)
val stmt1 = AssignStmt("x", e1) // x = y + z + exp(x - y)
// Let us make up the parts of the while statement
val cond1 = Leq(Ident("y"), Const(15.0f)) // y <= 15
// The assignment inside the while
val stmt2 = AssignStmt("y", Minus(Ident("y"), Ident("x"))) // y = y - x
// The condition for the if statement
val cond2 = Leq(Ident("x"), Const(0.0f)) // x <= 0.0
val stmt3 = AssignStmt("x", Minus(Const(0.0f), Ident("x"))) // x= 0 - x
val ifStmt = IfThenElseStmt(cond2, List(stmt3), List()) // if ( x <= 0) x = -x else nothing
val whileStmt = WhileStmt(cond1, List(stmt2, ifStmt)) // while loop 
val allstmts = List(stmt1, whileStmt)
val retExpr = Minus(Ident("y"), Ident("x")) // y - x to be returned

val c2 = DefaultCase(List(ifStmt))

passed(5)
//END TEST
```

<!-- grade_id: cell-73f9de24de44bb2f, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Plus(Ident("y"), Plus( Ident("z"), Exp( Minus(Ident("x"), Ident("y")) ) )) // y + z + exp(x - y)
val stmt1 = AssignStmt("x", e1) // x = y + z + exp(x - y)
// Let us make up the parts of the while statement
val cond1 = Leq(Ident("y"), Const(15.0f)) // y <= 15
// The assignment inside the while
val stmt2 = AssignStmt("y", Minus(Ident("y"), Ident("x"))) // y = y - x
// The condition for the if statement
val cond2 = Leq(Ident("x"), Const(0.0f)) // x <= 0.0
val stmt3 = AssignStmt("x", Minus(Const(0.0f), Ident("x"))) // x= 0 - x
val ifStmt = IfThenElseStmt(cond2, List(stmt3), List()) // if ( x <= 0) x = -x else nothing
val whileStmt = WhileStmt(cond1, List(stmt2, ifStmt)) // while loop 
val allstmts = List(stmt1, whileStmt)
val retExpr = Minus(Ident("y"), Ident("x")) // y - x to be returned

val c1 = SwitchCase(Const(2.0), List(stmt1, stmt3))
val c3 = SwitchCase(Ident("y"), List(stmt1, stmt2, whileStmt))
val c2 = DefaultCase(List(ifStmt))

val v1 = Switch(e1, List(c1, c2, c3))

passed(5)
//END TEST
```

## Problem 2 (15 Points)

Add list manipulation expressions to the allowed expressions in the `mython` language. 

We allow list expressions (use a new non terminal called __ListExpr__) that can do the following:
- The empty list: use constructor symbol `EmptyList`. It should not have arguments.
- The new list of arithmetic expressions: use constructor symbol `NewList`. It has as arguments a list of arithmetic expressions. Use Kleene-star in your grammar to specify a list.
- "Cons" an arithmetic expression to the front of a list: use constructor symbol `Cons`. It must have two arguments: the first argument is an arithmetic expression and the second argument is a list expression.
- Concatenate two lists defined by list expressions: use constructor symbol `Concat`. It must have two arguments that are list expressions.

Next, we add the following rule to  __CondExpr__  (conditional expressions):
- A condition expression to check if a list is empty:  Use constructor `IsEmptyList`. It's argument must be a list expression that we check for emptiness.

Finally, we add a special foreach loop statement over the elements of a list.

`foreach (x in list_expr) begin  .. list of statements...  end `

- Add a statement to loop over the elements of a list: use constructor`ForEach`. Its arguments must include (a) the identifier, (b) the list expression to be iterated over and (c) the body of the loop (list of statements).

Write down the extra grammar rules that you would add to the overall 
`mython` grammar to create these functionalities.

<!-- grade_id: cell-76cfd1d2c6a5898d, points: 15, grade: true, solution: true -->
YOUR ANSWER HERE

## Problem 3 (20 points)

### A (5 points)
We defined lists in the class. Write a recursive procedure to get the nth element of the list or throw an `IllegalArgumentException` if $n < 0$ or $n >= \text{length of list} $. Assume $n=0$ obtains the very first element
and $n = \text{length of list} -1$ yields very last element. 

<!-- grade_id: cell-724b4ca89445ceaf, points: 0, grade: false, solution: true -->
```scala
sealed trait NumList
case object Nil extends NumList
case class Cons(n: Int, l: NumList) extends NumList 

def getNthElement(lst: NumList, n: Int): Int = {
    ??? // YOUR CODE HERE
}
```

<!-- grade_id: cell-2bc6d5ee94793b48, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val l1 = Nil
val l2 = Cons(1, Cons(-1, Nil))
val l3 = Cons(1, Cons(2, l2))
val l4 = Cons(0, Cons(4, Cons(8, l3)))

val test1 = try {
    getNthElement(Nil, 3);
    assert(false, "Test 1 : getNthElement(Nil, 3) should raise an IllegalArgumentException. Your code did not.")
} catch {
    case e: IllegalArgumentException => "OK"
} 

assert(getNthElement(l2, 0) == 1, "Test2: getNthElement(l2, 0)  failed (expected answer = 1)")
assert(getNthElement(l3, 3) == -1, "Test3: getNthElement(l3, 3)  failed (expected answer = -1)")
assert(getNthElement(l4, 2) == 8, "Test4: getNthElement(l4, 2)  failed (expected answer = 8)")

val test2 = try {
    getNthElement(l4, 8);
    assert(false, "Test 5 : getNthElement(l4, 8) should raise an IllegalArgumentException. Your code did not.")
} catch {
    case e: IllegalArgumentException => "OK"
}

passed(5)
//END TEST
```

### B (7 points)
Write a recursive procedure that returns true if the list has the Fibonacci property. I.e, every element at position $i \geq 2$ is the sum of the two preceding elements. Note that the property is trivially true for lists of sizes $0$ and $1$.

<!-- grade_id: cell-eedd55e1f86bd6c6, points: 0, grade: false, solution: true -->
```scala
def isFibonacciList(lst: NumList): Boolean = {
    ??? // YOUR CODE HERE
}
```

<!-- grade_id: cell-33a4f21c01f0c7c4, points: 7, grade: true, solution: false -->
```scala
//BEGIN TEST
val l1 = Cons(12, Cons(25, Cons(37, Nil)))
assert(isFibonacciList(l1), 
       "Test case 1 :  isFibonacciList(l1) -- should return true")

val l2 = Cons(14, Cons(-1, Cons(13, l1 )))
assert(isFibonacciList(l2), 
       "Test case 2 :  isFibonacciList(l2) -- should return true")

val l3 = Cons(7, Cons(7, l2))

assert(!isFibonacciList(l3), 
       "Test case 3 :  isFibonacciList(l3) -- should return false")

val l4 = Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Nil))))))
assert(isFibonacciList(l2), 
       "Test case 4:  isFibonacciList(l4) -- should return true")

passed(7)
//END TEST
```

### C (8 points)
Write a recursive function `filterNumList(l: NumList, f: Int => Boolean): NumList` that takes in a `NumList` and a function `f: Int => Boolean`. 

1. It should return a new list that consist of all elements of the list `l` that return `true` when the function `f` is called on them.
2. The returned list elements must preserve the same order as in the original list.

<!-- grade_id: cell-0f9f5a51ae34b3bb, points: 0, grade: false, solution: true -->
```scala
??? // YOUR CODE HERE
```

<!-- grade_id: cell-a2549ef77f3f9bb0, points: 8, grade: true, solution: false -->
```scala
//BEGIN TEST
val l1 = Cons(12, Cons(25, Cons(37, Nil)))
def f1(j: Int): Boolean =  j <= 25 && j >= 12
assert(filterNumList(l1, f1) == Cons(12, Cons(25, Nil)), "Test 1 failed.")


val l2 = Cons(22, Cons(135, Cons(137, l1)))
def f2(j: Int): Boolean =  j % 5 == 0
assert(filterNumList(l2, f2) == Cons(135, Cons(25, Nil)), "Test 2 failed.")

def f3(j: Int): Boolean =  j >= 210
assert(filterNumList(l2, f3) == Nil, "Test 3 failed.")
assert(filterNumList(Nil, f3) == Nil, "Test 4 failed.")

val l4 = Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Nil))))))

def f4(j: Int): Boolean =  j <= 0
assert(filterNumList(l4, f4) == l4, "Test 5 failed")

passed(8)
//END TEST
```