```python
# Initialize Otter
import otter
grader = otter.Notebook("nb2.ipynb")
```

# nb02:  Basics with Pandas and And In-Depth with NumPy
***

In this notebook we'll explore both Pandas and NumPy.  We'll start with some basic dataframe manipulation using Pandas.  Since Pandas is built on the backbone of NumPy we'll then dive into more NumPy functionality.


---
### Pandas module:

[Pandas](https://pandas.pydata.org/) is one of the most widely used Python libraries in data science. In this lab, you will review commonly used data wrangling operations/tools in Pandas. We aim to give you familiarity with:

* Creating DataFrames
* Slicing DataFrames (i.e. selecting rows and columns)
* Filtering data (using boolean arrays and groupby.filter)
* Aggregating (using groupby.agg)

In this nb you are going to use several pandas methods. Reminder from lecture that you may press `shift+tab` on method parameters to see the documentation for that method. For example, if you were using the `drop` method in pandas, you could press shift+tab to see what `drop` is expecting.



**Note**: The Pandas interface is notoriously confusing for beginners, and the documentation is not consistently great. Throughout the semester, you will have to search through Pandas documentation and experiment, but remember it is part of the learning experience and will help shape you as a data scientist!

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
import plotly.express as px
%matplotlib inline
```

## Creating DataFrames & Basic Manipulations

Recall that a [DataFrame](https://pandas.pydata.org/pandas-docs/stable/user_guide/dsintro.html#dataframe) is a table in which each column has a specific data type; there is an index over the columns (typically string labels) and an index over the rows (typically ordinal numbers).

Usually you'll create DataFrames by using a function like `pd.read_csv`. However, in this section, we'll discuss how to create them from scratch.

The [documentation](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.html) for the pandas `DataFrame` class provides several constructors for the DataFrame class.

**Syntax 1:** You can create a DataFrame by specifying the columns and values using a dictionary as shown below. 

The keys of the dictionary are the column names, and the values of the dictionary are lists containing the row entries.

```python
fruit_info = pd.DataFrame(
    data = {'fruit': ['apple', 'orange', 'banana', 'raspberry'],
          'color': ['red', 'orange', 'yellow', 'pink'],
          'price': [1.0, 0.75, 0.35, 0.05]
          })
fruit_info
```

Output:
```
fruit   color  price
0      apple     red   1.00
1     orange  orange   0.75
2     banana  yellow   0.35
3  raspberry    pink   0.05
```

**Syntax 2:** You can also define a DataFrame by specifying the rows as shown below. 

Each row corresponds to a distinct tuple, and the columns are specified separately.

```python
fruit_info2 = pd.DataFrame(
    [("red", "apple", 1.0), ("orange", "orange", 0.75), ("yellow", "banana", 0.35),
     ("pink", "raspberry", 0.05)], 
    columns = ["color", "fruit", "price"])
fruit_info2
```

Output:
```
color      fruit  price
0     red      apple   1.00
1  orange     orange   0.75
2  yellow     banana   0.35
3    pink  raspberry   0.05
```

You can obtain the dimensions of a DataFrame by using the shape attribute `DataFrame.shape`.

```python
fruit_info.shape
```

Output:
```
(4, 3)
```

You can also convert the entire DataFrame into a two-dimensional NumPy array.

```python
fruit_info.values
```

Output:
```
array([['apple', 'red', 1.0],
       ['orange', 'orange', 0.75],
       ['banana', 'yellow', 0.35],
       ['raspberry', 'pink', 0.05]], dtype=object)
```

### **REVIEW:** Selecting Rows and Columns in Pandas

As you've seen in lecture, there are two verbose operators in Python for selecting rows: `loc` and `iloc`. Let's review them briefly.

#### Approach 1: `loc`

The first of the two verbose operators is `loc`, which takes two arguments. The first is one or more row **labels**, the second is one or more column **labels** - both of which are displayed in bold to the left of each of the rows and above each of the columns respectively. These are not the same as positional indices, which are used for indexing Python lists or NumPy arrays!

The desired rows or columns can be provided individually, in slice notation, or as a list. Some examples are given below.

Note that **slicing in `loc` is inclusive** on the provided labels.

```python
#get rows 0 through 2 and columns fruit through price
fruit_info.loc[0:2, 'fruit':'price']
```

Output:
```
fruit   color  price
0   apple     red   1.00
1  orange  orange   0.75
2  banana  yellow   0.35
```

```python
# get rows 0 through 2 and columns fruit and price. 
# Note the difference in notation and result from the previous example.
fruit_info.loc[0:2, ['fruit', 'price']]
```

Output:
```
fruit  price
0   apple   1.00
1  orange   0.75
2  banana   0.35
```

```python
# get rows 0 and 2 and columns fruit and price. 
fruit_info.loc[[0, 2], ['fruit', 'price']]
```

Output:
```
fruit  price
0   apple   1.00
2  banana   0.35
```

```python
# get rows 0 and 2 and column fruit
fruit_info.loc[[0, 2], ['fruit']]
```

Output:
```
fruit
0   apple
2  banana
```

Note that if we request a single column but don't enclose it in a list, the return type of the `loc` operator is a `Series` rather than a DataFrame. 

```python
# get rows 0 and 2 and column fruit, returning the result as a Series
fruit_info.loc[[0, 2], 'fruit']
```

Output:
```
0     apple
2    banana
Name: fruit, dtype: object
```

If we provide only one argument to `loc`, it uses the provided argument to select rows, and returns all columns.

```python
fruit_info.loc[0:1]
```

Output:
```
fruit   color  price
0   apple     red   1.00
1  orange  orange   0.75
```

Note that if you try to access columns without providing rows, `loc` will crash. 

```python
# uncomment, this code will crash
#fruit_info.loc[["fruit", "price"]]

# uncomment, this code works fine: 
#fruit_info.loc[:, ["fruit", "price"]]
```

#### Approach 2: `iloc`

`iloc` is very similar to `loc` except that its arguments are row numbers and column numbers, rather than row labels and labels names. A usueful mnemonic is that the `i` stands for "integer". This is quite similar to indexing into a Python list or NumPy array.

In addition, **slicing for `iloc` is exclusive** on the provided integer indices. Some examples are given below:

```python
# get rows 0 through 3 (exclusive) and columns 0 through 2 (exclusive)
fruit_info.iloc[0:3, 0:3]
```

Output:
```
fruit   color  price
0   apple     red   1.00
1  orange  orange   0.75
2  banana  yellow   0.35
```

```python
# get rows 0 through 3 (exclusive) and columns 0 and 2.
fruit_info.iloc[0:3, [0, 2]]
```

Output:
```
fruit  price
0   apple   1.00
1  orange   0.75
2  banana   0.35
```

```python
# get rows 0 and 2 and columns 0 and 2.
fruit_info.iloc[[0, 2], [0, 2]]
```

Output:
```
fruit  price
0   apple   1.00
2  banana   0.35
```

```python
#get rows 0 and 2 and column fruit
fruit_info.iloc[[0, 2], [0]]
```

Output:
```
fruit
0   apple
2  banana
```

```python
# get rows 0 and 2 and column fruit
fruit_info.iloc[[0, 2], 0]
```

Output:
```
0     apple
2    banana
Name: fruit, dtype: object
```

Note that in these loc and iloc examples above, the row **label** and row **number** were always the same.

Let's see an example where they are different. If we sort our fruits by color, we get:

```python
fruit_info_sorted = fruit_info.sort_values("price")
fruit_info_sorted
```

Output:
```
fruit   color  price
3  raspberry    pink   0.05
2     banana  yellow   0.35
1     orange  orange   0.75
0      apple     red   1.00
```

Observe that the row number 0 now has index 3, row number 1 now has index 2, etc. These indices are the arbitrary numerical index generated when we created the DataFrame. For example, banana was originally in row 2, and so it has row label 2.

If we request the rows in positions 0 and 2 using `iloc`, we're indexing using the row NUMBERS, not labels. 

```python
fruit_info_sorted.iloc[[0, 2], 0]
```

Output:
```
3    raspberry
1       orange
Name: fruit, dtype: object
```

Lastly, similar to with `loc`, the second argument to `iloc` is optional. That is, if you provide only one argument to `iloc`, it treats the argument you provide as a set of desired row numbers, not column numbers.

```python
fruit_info.iloc[[0, 2]]
```

Output:
```
fruit   color  price
0   apple     red   1.00
2  banana  yellow   0.35
```

#### Approach 3: `[]` Notation for Accessing Rows and Columns

Pandas also supports a bare `[]` operator. It's similar to `loc` in that it lets you access rows and columns by their name.

However, unlike `loc`, which takes row names and also optionally column names, `[]` is more flexible. If you provde it only row names, it'll give you rows (same behavior as `loc`), and if you provide it with only column names, it'll give you columns (whereas `loc` will crash).

Some examples:

```python
fruit_info[0:2]
```

Output:
```
fruit   color  price
0   apple     red   1.00
1  orange  orange   0.75
```

```python
# Here we're providing a list of fruits as single argument to []
fruit_info[["fruit", "color", "price"]]
```

Output:
```
fruit   color  price
0      apple     red   1.00
1     orange  orange   0.75
2     banana  yellow   0.35
3  raspberry    pink   0.05
```

Note that slicing notation is not supported for columns if you use `[]` notation. Use `loc` instead.

```python
# uncomment and this code crashes
#fruit_info["fruit":"price"]

# uncomment and this works fine
#fruit_info.loc[:, "fruit":"price"]
```

Output:
```
File /opt/conda/lib/python3.10/site-packages/pandas/core/indexes/base.py:6373, in Index._validate_indexer(self, form, key, kind)
   6370 assert kind in ["getitem", "iloc"]
   6372 if key is not None and not is_integer(key):
-> 6373     self._raise_invalid_indexer(form, key)

File /opt/conda/lib/python3.10/site-packages/pandas/core/indexes/base.py:4151, in Index._raise_invalid_indexer(self, form, key, reraise)
   4149 if reraise is not lib.no_default:
   4150     raise TypeError(msg) from reraise
-> 4151 raise TypeError(msg)

TypeError: cannot do slice indexing on RangeIndex with these indexers [fruit] of type str
```

`[]` and `loc` are quite similar. For example, the following two pieces of code are functionally equivalent for selecting the fruit and price columns.

1. `fruit_info[["fruit", "price"]]` 
2. `fruit_info.loc[:, ["fruit", "price"]]`.

Because it yields more concise code, you'll find that our code and your code both tend to feature `[]`. However, there are some subtle pitfalls of using `[]`. If you're ever having performance issues, weird behavior, or you see a `SettingWithCopyWarning` in pandas, switch from `[]` to `loc` and this may help.

To avoid getting too bogged down in indexing syntax, we'll avoid a more thorough discussion of `[]` and `loc`. We may return to this at a later point in the course.

For more on `[]` vs `loc`, you may optionally try reading:
1. https://stackoverflow.com/questions/48409128/what-is-the-difference-between-using-loc-and-using-just-square-brackets-to-filte
2. https://stackoverflow.com/questions/38886080/python-pandas-series-why-use-loc/65875826#65875826
3. https://stackoverflow.com/questions/20625582/how-to-deal-with-settingwithcopywarning-in-pandas/53954986#53954986

Now that we've reviewed basic indexing, let's discuss how we can modify dataframes. We'll do this via a series of exercises. 

### Question 1(a)

For a DataFrame `d`, you can add a column with `d['new column name'] = ...` and assign a list or array of values to the column. Add a column of integers containing 1, 2, 3, and 4 called `rank1` to the `fruit_info` table which expresses your personal preference about the taste ordering for each fruit (1 is tastiest; 4 is least tasty). 


<!-- grade_id: q1_answer, points: 0, grade: false, solution: true -->
```python
fruit_info['rank1'] = [1, 2, 3, 4]
```

```python
print(fruit_info["rank1"].dtype)
```

Output:
```
int64
```

```python
grader.check("q1a")
```

Output:
```
q1a results:
    q1a - 1 result:
        ❌ Test case failed
        Trying:
            assert fruit_info["rank1"].dtype == 'int64'
        Expecting:
            True
        **********************************************************************
        Line 1, in q1a 0
        Failed example:
            assert fruit_info["rank1"].dtype == 'int64'
        Expected:
            True
        Got nothing

    q1a - 2 result:
        ✅ Test case passed
```

### Question 1(b)

You can also add a column to `d` with `d.loc[:, 'new column name'] = ...`. As above, the first parameter is for the rows and second is for columns. The `:` means change all rows and the `'new column name'` indicates the name of the column you are modifying (or in this case, adding). 

Add a column called `rank2` to the `fruit_info` table which contains the same values in the same order as the `rank1` column.


<!-- grade_id: cell-ef625b2f6154e9b2, points: 0, grade: false, solution: true -->
```python
fruit_info.loc[:,'rank2'] = fruit_info['rank1']
```

```python
grader.check("q1b")
```

Output:
```
q1b results: All test cases passed!
```

### Question 2

Use the `.drop()` method to [drop](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.drop.html) both the `rank1` and `rank2` columns you created. Make sure to use the `axis` parameter correctly. Note that `drop` does not change a table, but instead returns a new table with fewer columns or rows unless you set the optional `inplace` parameter.

*Hint*: Look through the [documentation](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.drop.html) to see how you can drop multiple columns of a Pandas DataFrame at once using a list of column names.


<!-- grade_id: q2_answer, points: 0, grade: false, solution: true -->
```python
fruit_info_original = fruit_info.drop(labels=['rank1', 'rank2'], axis=1)
fruit_info_original
```

Output:
```
fruit   color  price
0      apple     red   1.00
1     orange  orange   0.75
2     banana  yellow   0.35
3  raspberry    pink   0.05
```

```python
print(fruit_info_original.shape)
assert fruit_info_original.shape == (4,3)
```

Output:
```
(4, 3)
```

```python
print(fruit_info.shape)
```

Output:
```
(4, 5)
```

```python
grader.check("q2")
```

Output:
```
q2 results:
    q2 - 1 result:
        ❌ Test case failed
        Trying:
            assert fruit_info_original.shape == (4,3)
        Expecting:
            True
        **********************************************************************
        Line 1, in q2 0
        Failed example:
            assert fruit_info_original.shape == (4,3)
        Expected:
            True
        Got nothing

    q2 - 2 result:
        ❌ Test case failed
        Trying:
            assert fruit_info.shape == (4,5)
        Expecting:
            (4, 5)
        **********************************************************************
        Line 1, in q2 1
        Failed example:
            assert fruit_info.shape == (4,5)
        Expected:
            (4, 5)
        Got nothing

    q2 - 3 result:
        ✅ Test case passed
```

### Question 3

Use the `.rename()` method to [rename](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.rename.html) the columns of `fruit_info_original` so they begin with capital letters. Set this new DataFrame to `fruit_info_caps`. For an example of how to use rename, see the linked documentation above.

<!-- grade_id: q3_answer, points: 0, grade: false, solution: true -->
```python
fruit_info_caps = fruit_info_original.rename(columns={'fruit':'Fruit', 'color':'Color', 'price':'Price'})
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

---


## Switching Gears:  Deep Diving Into Numpy 


### NumPy module:

𝐍𝐮𝐦𝐏𝐲  (short for "Numeric Python") is an open source Python module for scientific computing. NumPy supports large, multidimensional arrays and matrices. It also supports a large collection of mathematical functions not found in Python's standard math library.

It's custom in data science to import Numpy  with the alias $\texttt{np}$.  We can then ***access any mathematical function in the Numpy library by prepending function names by $\texttt{np}$***

For a more in-depth description of methods and attributes built into Numpy see here:  https://numpy.org/doc/stable/user/quickstart.html   
and  here:  https://numpy.org/doc/stable/reference/routines.math.html



### Mathematical Functions

Numpy includes all of the functions and mathematical constants in Python's standard math library, like logarithms, exponentiation, and even $\pi$.  

```python
print(np.log(np.exp(1)))
print(np.log2(16))
print(np.log10(1000))
print(np.pi)
```

Output:
```
0.9999999999999999
4.0
3.0
3.141592653589793
```

The nice thing about Numpy's mathematical functions is that they can be applied to arrays as well as scalars.  For instance

```python
u = np.array([10, 100, 1000, 10000])

np.log10(u)
```

Output:
```
array([1., 2., 3., 4.])
```

### Question 4

Use numpy built-in functions to find the square root of 17: </span>


```python
#Exercise 4 
np.sqrt(17)
```

Output:
```
4.123105625617661
```

```python
# if you don't like how cumbersome it is to use np as the prefix, there is a solution!
from numpy import pi

print(pi)

#Or, to practice printing this with specific formatting:

"Floating point pi = {0:.3f}, with {1:d} digit precision".format(pi, 3)
#Here we specify 3 digits of precision and f is used to represent floating point number
#d in {1:d} represents integer value.
#See this site for many ways to format your output: https://thepythonguru.com/python-string-formatting/
```

Output:
```
3.141592653589793
'Floating point pi = 3.142, with 3 digit precision'
```

## The Numpy ndarray
The main workhorse of Numpy is its  multidimensional array  class called the  𝚗𝚍𝚊𝚛𝚛𝚊𝚢 . 

A numpy array is a grid of values that belong to a similar data type.

The numpy array values are indexed by a tuple of nonnegative integers.

This allows us to build and work with one dimensional, two dimensional, or even  𝑁 -dimensional arrays of numbers.

We can build a  one dimensional  ndarray by passing a Python list to the  𝚗𝚙.𝚊𝚛𝚛𝚊𝚢()  function.

See here for a full list of the array attributes and methods:  https://docs.scipy.org/doc/numpy-1.6.0/reference/generated/numpy.ndarray.html

### Creating Arrays

You can build arrays from python lists. 


```python
np.array([[1.,2.], [3.,4.]])
```

Output:
```
array([[1., 2.],
       [3., 4.]])
```

```python
np.array([x for x in range(5)])
```

Output:
```
array([0, 1, 2, 3, 4])
```

Array's don't have to contain numbers:

```python
np.array([["A", "matrix"], ["of", "words."]])
```

Output:
```
array([['A', 'matrix'],
       ['of', 'words.']], dtype='<U6')
```

## Making Arrays of Zeros

```python
np.zeros(5)
```

Output:
```
array([0., 0., 0., 0., 0.])
```

## Making Arrays of Ones

```python
np.ones([3,2])
```

Output:
```
array([[1., 1.],
       [1., 1.],
       [1., 1.]])
```

```python
np.eye(4)
```

Output:
```
array([[1., 0., 0., 0.],
       [0., 1., 0., 0.],
       [0., 0., 1., 0.],
       [0., 0., 0., 1.]])
```

## Making Arrays from ranges:

The `np.arange(start, stop, step)` function is like the python `range` function.

```python
np.arange(0, 10, 2)
```

Output:
```
array([0, 2, 4, 6, 8])
```

You can make a range of other types as well:

```python
np.arange(np.datetime64('2016-12-31'), np.datetime64('2017-02-01'))
```

Output:
```
array(['2016-12-31', '2017-01-01', '2017-01-02', '2017-01-03',
       '2017-01-04', '2017-01-05', '2017-01-06', '2017-01-07',
       '2017-01-08', '2017-01-09', '2017-01-10', '2017-01-11',
       '2017-01-12', '2017-01-13', '2017-01-14', '2017-01-15',
       '2017-01-16', '2017-01-17', '2017-01-18', '2017-01-19',
       '2017-01-20', '2017-01-21', '2017-01-22', '2017-01-23',
       '2017-01-24', '2017-01-25', '2017-01-26', '2017-01-27',
       '2017-01-28', '2017-01-29', '2017-01-30', '2017-01-31'],
      dtype='datetime64[D]')
```

## Interpolating numbers 

The `linspace(start,end,num)` function generates `num` numbers evenly spaced between the `start` and `end`.

```python
np.linspace(0, 5, 10)
```

Output:
```
array([0.        , 0.55555556, 1.11111111, 1.66666667, 2.22222222,
       2.77777778, 3.33333333, 3.88888889, 4.44444444, 5.        ])
```

Learn more about working with [datetime objects](https://docs.scipy.org/doc/numpy/reference/arrays.datetime.html#).

## A random array 

You can also generate arrays of random numbers (we will cover this in greater detail later).


- `rand` generates random numbers from a Uniform(low=0, high=1) distribution.
- `permutation` generates a random permutation of a sequence of numbers.

```python
np.random.rand(3,2)
```

Output:
```
array([[0.62043432, 0.52914404],
       [0.84551248, 0.27235825],
       [0.06798253, 0.11400161]])
```

```python
np.random.permutation(range(0,10))
```

Output:
```
array([2, 5, 3, 9, 1, 8, 4, 7, 0, 6])
```

```python
# generate random integers
np.random.randint(0,20,5) # low, high, size
```

Output:
```
array([15,  6, 18,  1, 10])
```


# Properties of Arrays

## Shape

Arrays have a shape which corresponds to the number of rows, columns, fibers, ...

```python
A = np.array([[1., 2., 3.], [4., 5., 6.]])
print(A)
A.shape
```

Output:
```
[[1. 2. 3.]
 [4. 5. 6.]]
(2, 3)
```

## Type

Arrays have a type which corresponds to the type of data they contain

```python
A.dtype
```

Output:
```
dtype('float64')
```

```python
np.arange(1,5).dtype
```

Output:
```
dtype('int64')
```

```python
(np.array([True, False])).dtype
```

Output:
```
dtype('bool')
```

```python
np.array(["Hello", "Worlddddd!"]).dtype
```

Output:
```
dtype('<U10')
```

What does `<U6` mean?

- `<` Little Endian
- `U` Unicode
- `6` length of longest string

#### and we can change the type of an array:

```python
np.array([1,2,3]).astype(float)
```

Output:
```
array([1., 2., 3.])
```

```python
np.array(["1","2","3"]).astype(int)
```

Output:
```
array([1, 2, 3])
```

Learn more about numpy [array types](https://docs.scipy.org/doc/numpy/user/basics.types.html)

```python
# 'y' is a list in Python
y = [1,2,3,4,5]
print(y)

#We can check the type of what we've just created with Python's type function.
type(y)
```

Output:
```
[1, 2, 3, 4, 5]
list
```

```python
#Now, the list is an array
x = np.array(y, dtype=float)
print(x)

type(x)
```

Output:
```
[1. 2. 3. 4. 5.]
numpy.ndarray
```

```python
# The attribute .dtype will let you see the datatype of the elements: 

print(x.dtype)
```

Output:
```
float64
```

### Question 5
    
Write code to create a vector of zeros of size 10 and update the sixth entry to 11 

```python
ten_zeros = np.zeros(10)
ten_zeros[5] = 11
print(ten_zeros)
```

Output:
```
[ 0.  0.  0.  0.  0. 11.  0.  0.  0.  0.]
```

```python
ten_zeros.dtype
```

Output:
```
dtype('float64')
```

```python
#We are not limited to 1-dimensional arrays - we can also create  n-Dimensional arrays!
# 1-dimensional array
z=np.array([14, 15, 22.5, 100])
# 2-dimensional array
w=np.array([[1, 2],
       [3, 4],
       [5, 6]])
```

<img src="images/NumpyArrays.png"/>

```python
#You can find the dimension(s) of your array using the shape attribute:

print(z.shape)

print(w.shape)
```

Output:
```
(4,)
(3, 2)
```

The nice thing about using Numpy arrays to store numerical values is that we can perform mathematical operations on them.  For instance, if we want to multiply every value in $\texttt{x}$ by $5$ and store the result in a new array $\texttt{y}$ we can do so as follows: 

```python
y = 5*x
print(y)
```

Output:
```
[ 5. 10. 15. 20. 25.]
```

Contrast this to what happens if we tried the exact same operations on a list instead:


```python
w = [1, 2, 3, 4, 5]
# w is a Python list

q = 5 * w
# q is a list of lists
print('w = ', w)
print('5 * w = ', q)
```

Output:
```
w =  [1, 2, 3, 4, 5]
5 * w =  [1, 2, 3, 4, 5, 1, 2, 3, 4, 5, 1, 2, 3, 4, 5, 1, 2, 3, 4, 5, 1, 2, 3, 4, 5]
```

If we want to add two ndarray's of equal length together we can do that too: 

```python
z = x + y 
print(z)
```

Output:
```
[ 6. 12. 18. 24. 30.]
```

If we want to create a two dimensional array we simply pass a list of lists to the $\texttt{np.array()}$ function.  The lists in the list of lists then become the rows of the two dimensional array. 

```python
A = np.array([[1,2,3,4,5], [6,7,8,9,10]], dtype=float)
print(A)
```

Output:
```
[[ 1.  2.  3.  4.  5.]
 [ 6.  7.  8.  9. 10.]]
```

We can access elements of Numpy arrays in a ways similar to the way we access elements in Python lists. For instance, if we want to get the first $3$ elements of the array $\texttt{y}$, we can do so as follows: 

```python
print(y)
```

Output:
```
[ 5. 10. 15. 20. 25.]
```

```python
y[0:4]
```

Output:
```
array([ 5., 10., 15., 20.])
```

Just like with Python lists, if we're indexing from the start of the array there is no need to include the $0$ in the index range.  We can simply do 

```python
y[:3]
```

Output:
```
array([ 5., 10., 15.])
```

Similarly, if we want to access everything from the third entry to the end of the array, we can do 

```python
y[2:]
```

Output:
```
array([15., 20., 25.])
```

Indexing in multidimensional arrays is similar.  First, let's build a two dimensional array with our arrays $\texttt{x}, \texttt{y},$ and $\texttt{z}$ as the rows. 

```python
B = np.array([x,y,z])
print(B)
```

Output:
```
[[ 1.  2.  3.  4.  5.]
 [ 5. 10. 15. 20. 25.]
 [ 6. 12. 18. 24. 30.]]
```

We can slice up the two dimensional array by doing slices along rows and columns.  Let's suppose we wanted to carve out the second row of $\texttt{B}$. 

```python
B[1,:]

```

Output:
```
array([ 5., 10., 15., 20., 25.])
```

Here, the stuff before the comma refers to rows of the array, and the stuff after the comma refers to columns of the array.  In the previous command we've indicated that we want the row with index $1$, and the colon indicated that we want _all_ of the columns. Similarly, if we want all of the fourth column, we could do 

```python
B[:,3]
```

Output:
```
array([ 4., 20., 24.])
```

If you've played with slicing lists of Python lists then you'll note that the main difference between slicing two dimensional Numpy arrays is that we only need one set of brackets, instead of two. Otherwise it's pretty similar.   

If we want to carve out certain rows and certain columns, we can do that by grabbing ranges of indices for both the rows and the columns.  For instance, if we want the part of $\texttt{B}$ in the second and third rows and the second through the fourth columns we can do 

```python
B[1:3, 1:4]
```

Output:
```
array([[10., 15., 20.],
       [12., 18., 24.]])
```

We can also apply mathematical functions to the two dimensional array.  For instance, if we want to sum all of the entries in $\texttt{B}$ we can do 

```python
np.sum(B)
```

Output:
```
180.0
```

If instead we just want to sum along the rows or columns of the array we can add the $\texttt{axis}$ parameter.  

```python
np.sum(B, axis=0)
```

Output:
```
array([12., 24., 36., 48., 60.])
```

Notice that choosing $\texttt{axis=0}$ caused us to sum along the columns of $\texttt{B}$.  If instead we used $\texttt{axis=1}$ we would sum along the rows 

```python
np.sum(B, axis=1)
```

Output:
```
array([15., 75., 90.])
```

We can also index into an array by conditions.  For instance, remember the vector $\texttt{z}$

```python
print(z)
```

Output:
```
[ 6. 12. 18. 24. 30.]
```

Let's suppose I want to grab all of the entries in $\texttt{z}$ that are bigger than $15$.  The condition $\texttt{z > 15}$ returns a boolean array indicating whether each entry in $\texttt{z}$ satisfies the given condition 

```python
z > 15
```

Output:
```
array([False, False,  True,  True,  True])
```

Now, if I want to actually extract those entries of $\texttt{z}$ that satisfy the condition we can index $\texttt{z}$ using this boolean array 

```python
z[z>15]

```

Output:
```
array([18., 24., 30.])
```

### Question 6
Use numpy built-in functions to create an array with elements [2, 4, 8] and then output an array with values $[e^2, e^4, e^8]$ 

```python
pow2 = np.array([2, 4, 8])
eees = np.array([np.e, np.e, np.e])

eees**pow2
```

Output:
```
array([   7.3890561 ,   54.59815003, 2980.95798704])
```

### Accessing elements in Arrays

We can access elements of Numpy arrays in a way similar to the way we access elements in Python lists. For instance, if we want to get the first $3$ elements of the array $\texttt{y}$, we can do so as follows: 

```python
print(y)
# recall y is an nparray

print(y[0:3])


```

Output:
```
[ 5. 10. 15. 20. 25.]
[ 5. 10. 15.]
```

Just like with Python lists, if we're indexing from the start of the array there is no need to include the $0$ in the index range.  We can simply do 

```python
y[:3]
```

Output:
```
array([ 5., 10., 15.])
```

Similarly, if we want to access everything from the third entry (indexed with a 2) to the end of the array, we can do 

```python
# Recall y = [5. 10. 15. 20. 25.]
y[2:]
```

Output:
```
array([15., 20., 25.])
```

We can also index into an array by conditions.  For instance, remember the vector $\texttt{z}$

```python
print(z)
```

Output:
```
[ 6. 12. 18. 24. 30.]
```

Let's suppose I want to grab all of the entries in $\texttt{z}$ that are bigger than $15$.  The condition $\texttt{z > 15}$ returns a boolean array indicating whether each entry in $\texttt{z}$ satisfies the given condition 

```python
z > 15
#type(z)
```

Output:
```
array([False, False,  True,  True,  True])
```

Now, if I want to actually extract those entries of $\texttt{z}$ that satisfy the condition we can index $\texttt{z}$ using this boolean array 

```python
print(z)
z[z>15]
```

Output:
```
[ 6. 12. 18. 24. 30.]
array([18., 24., 30.])
```

### Question 7
Write code to extract the entries of z that are divisible by 4: 

```python
#Question 7

z[z%4==0]

```

Output:
```
array([12., 24.])
```

# Reshaping

Often you will need to reshape matrices.  Suppose you have the following array:

```python
np.arange(1,13)
```

Output:
```
array([ 1,  2,  3,  4,  5,  6,  7,  8,  9, 10, 11, 12])
```

**What will the following produce:**

```python
np.arange(1,13).reshape(4,3)
```

**Option A:**

```python
array([[ 1,  2,  3],
       [ 4,  5,  6],
       [ 7,  8,  9],
       [10, 11, 12]])
```

**Option B:**

```python
array([[ 1,  5,  9],
       [ 2,  6, 10],
       [ 3,  7, 11],
       [ 4,  8, 12]])
```

**Solution**

```python
A = np.arange(1,13).reshape(4,3)
A
```

Output:
```
array([[ 1,  2,  3],
       [ 4,  5,  6],
       [ 7,  8,  9],
       [10, 11, 12]])
```

## Flattening Matrix

Flattening a matrix (higher dimensional array) produces a one dimensional array.

```python
A.flatten()
```

Output:
```
array([ 1,  2,  3,  4,  5,  6,  7,  8,  9, 10, 11, 12])
```

Understanding the slice syntax

```python
begin:end:stride
```


## Modifying a Slice

Suppose I wanted to make all entries in my matrix 0 in the top right corner as in (H) above.

```python
H = np.arange(1,13).reshape(4,3)
print("Before:\n", H)
```

Output:
```
Before:
 [[ 1  2  3]
 [ 4  5  6]
 [ 7  8  9]
 [10 11 12]]
```

```python
H[:2, 1:] = 0
print("After:\n", H)
```

Output:
```
After:
 [[ 1  0  0]
 [ 4  0  0]
 [ 7  8  9]
 [10 11 12]]
```

# Boolean Indexing

We can apply boolean operations to arrays.  This is essential when trying to select and modify individual elements.


**Question:** *Given the following definition of A:*

```python
[[   1.    2.    3.]
 [   4.    5. -999.]
 [   7.    8.    9.]
 [  10. -999. -999.]]
```

*what will the following output:*
```python
A > 3
```


- **Option A:**

```python
False
```

- **Option B:**

```python
array([[False, False, False],
       [ True,  True, False],
       [ True,  True,  True],
       [ True, False, False]], dtype=bool)
```

```python
A = np.array([[  1.,   2.,   3.],
       [  4.,   5.,   -999.0],
       [  7.,   8.,   9.],
       [ 10.,  -999.0,  -999.0]])

A > 3.
```

Output:
```
array([[False, False, False],
       [ True,  True, False],
       [ True,  True,  True],
       [ True, False, False]])
```

**Question:** *What will the following output*
```python
A = np.array([[   1.,    2.,    3.],
       [   4.,    5., -999.],
       [   7.,    8.,    9.],
       [  10., -999., -999.]])

A[A > 3]
```


- **Option A:**

```python
array([ 4,  7, 10,  5,  8, 11,  6,  9, 12])
```

- **Option B:**

```python
array([  4.,   5.,   7.,   8.,   9.,  10.])
```


- **Option C:**

```python
array([[  nan,   nan,  nan],
       [  4.,    5.,   nan],
       [  7.,    8.,   9.],
       [ 10.,    nan,  nan]])
```

```python
A = np.array([[  1.,   2.,   3.],
       [  4.,   5.,   -999.0],
       [  7.,   8.,   9.],
       [ 10.,  -999.0,  -999.0]])

A[A > 3] 
```

Output:
```
array([ 4.,  5.,  7.,  8.,  9., 10.])
```

**Question:** *Replace the -999.0 entries with `np.nan`.*

```python
array([[   1.,    2.,    3.],
       [   4.,    5., -999.],
       [   7.,    8.,    9.],
       [  10., -999., -999.]])
```

**Solution**

```python
A = np.array([[  1.,   2.,   3.],
       [  4.,   5.,   -999.0],
       [  7.,   8.,   9.],
       [ 10.,  -999.0,  -999.0]])
```

* Construct a boolean array that indicates where the value is 999.0:

```python
ind = (A == -999.0)
print(ind)
```

Output:
```
[[False False False]
 [False False  True]
 [False False False]
 [False  True  True]]
```

* Assign `np.nan` to all the `True` entires:

```python
A[ind] = np.nan
A
```

Output:
```
array([[ 1.,  2.,  3.],
       [ 4.,  5., nan],
       [ 7.,  8.,  9.],
       [10., nan, nan]])
```

**Question:** *What might -999.0 represent?  Why might I want to replace the -999.0 with a `np.nan`?  *

**Solution:** It could be safer in calculations.  For example when computing the mean of the transformed A we get:

```python
print(A)
np.mean(A)
```

Output:
```
[[ 1.  2.  3.]
 [ 4.  5. nan]
 [ 7.  8.  9.]
 [10. nan nan]]
nan
```

Perhaps instead we want:

```python
np.nanmean(A)
```

Output:
```
5.444444444444445
```

```python
help(np.nanmean)
```

Output:
```
Help on function nanmean in module numpy:

nanmean(a, axis=None, dtype=None, out=None, keepdims=<no value>, *, where=<no value>)
    Compute the arithmetic mean along the specified axis, ignoring NaNs.
    
    Returns the average of the array elements.  The average is taken over
    the flattened array by default, otherwise over the specified axis.
    `float64` intermediate and return values are used for integer inputs.
    
    For all-NaN slices, NaN is returned and a `RuntimeWarning` is raised.
    
    .. versionadded:: 1.8.0
    
    Parameters
    ----------
    a : array_like
        Array containing numbers whose mean is desired. If `a` is not an
        array, a conversion is attempted.
    axis : {int, tuple of int, None}, optional
        Axis or axes along which the means are computed. The default is to compute
        the mean of the flattened array.
    dtype : data-type, optional
        Type to use in computing the mean.  For integer inputs, the default
        is `fl
```

# More Complex Bit Logic

Often we will want to work with multiple different arrays at once and select subsets of entries from each array.  Consider the following example:

```python
names = np.array(["Joey", "Henry", "Joseph", 
                  "Jim", "Sam", "Deb", "Mike", 
                  "Bin", "Joe", "Andrew", "Bob"])

favorite_number = np.arange(len(names)) 
```

Suppose a subset of these people are staff members:

```python
staff = ["Joey",  "Deb", "Sam"]
```

### Question 8:
*How could we compute the sum of the staff members favorite numbers?*

One solution is to use for loops:

```python

```

Another solution would be to use the [np.in1d](https://docs.scipy.org/doc/numpy/reference/generated/numpy.in1d.html) function to determine which people are staff.

```python

```

Boolean indexing

```python

```

### Question:
*What does the following expression compute:*

```python
starts_with_j = np.char.startswith(names, "J")
starts_with_j[is_staff].mean()
```

```python

```

### Question:
*What does it mean to take the mean of an array of booleans?*

### Question
*What does the following expression compute:*

```python
favorite_number[starts_with_j & is_staff].sum()
```

```python

```

## Other Useful Operations

* [`choose()`](https://docs.scipy.org/doc/numpy-1.10.0/reference/generated/numpy.choose.html)
* [`where()`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.where.html)



# A Note on using Array operations

```python
data = np.random.rand(1000000)
```

Consider the following two programs.  


### Program A
```python
s = 0
c = 0
for x in data:
    if x > 0.5:
        s += x
        c += 1
result = s/c
```

### Program B
```python
result = data[data > 0.5].mean()
```

1. What do they do?
1. Which one is faster?
1. Which one is clearer?

---

<br/><br/><br/><br/><br/>

### Solution

```python

```

```python

```

## Important Points

Using the array abstractions instead of looping can often be:

1. Clearer
2. Faster

These are fundamental goals of abstraction. 

<br/>
<br/>
<br/>




## Be Careful with Floating Point Numbers

What is the value of the following:
$$
A - \exp \left( \log \left( A \right) \right)
$$

<br/>
<br/>
<br/>

**Solution:**


```python
A = np.arange(1., 13.).reshape(4,3)
print(A)

(A - np.exp(np.log(A)))
```

**What happened?!**

Floating point precision is not perfect and we are applying transcendental functions.

<br/><br/><br/><br/>
### A simpler examples

What is the value of the following expression:

```python
0.1 + 0.2 == 0.3
```

```python
0.1 + 0.2 == 0.3
```

```python
print(0.1 + 0.2)
```

For these situations consider using `np.isclose`:

```python
help(np.isclose)
```

<br/>
<br/>
<br/>
<br/>


## Aggregating along an axis 


### Grouping by row:

```python
A.sum(axis=0)
```

This is the same as:
```python
(nrow, ncols) = A.shape

s = np.zeros(ncols)

for i in range(nrows):
    s += A[i,:]

print(s)
```

### Grouping by col:

```python
A.sum(axis=1)
```

This is the same as:
```python
(nrows, ncols) = A.shape

s = np.zeros(nrows)

for i in range(ncols):
    s += A[:,i]

print(s)
```

## Other Functions to Checkout

* [`mean`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.mean.html) computes the mean 
* [`std`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.std.html) standard deviation
* [`var`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.var.html) variance

and [many more](https://docs.scipy.org/doc/numpy/reference/ufuncs.html#math-operations)

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1a (points: None)
```
>>> assert fruit_info["rank1"].dtype == 'int64'
True
```
```
>>> assert sorted(fruit_info["rank1"].dropna()) == [1,2,3,4]
```


Tests for q1b (points: None)
```
>>> assert fruit_info["rank2"].dtype == 'int64'
```
```
>>> assert (fruit_info["rank2"] == fruit_info["rank1"]).all()
```


Tests for q2 (points: None)
```
>>> assert fruit_info_original.shape == (4,3)
True
```
```
>>> assert fruit_info.shape == (4,5)
(4, 5)
```
```
>>> assert (fruit_info_original.columns == ["fruit", "color", "price"]).all()
```


Tests for q3 (points: None)
```
>>> assert (fruit_info_caps.columns == ["Fruit", "Color", "Price"]).all()
```
```
>>> assert (fruit_info.columns == ['fruit', 'color', 'price', 'rank1', 'rank2']).all()
```
